Sperm motility in mice with oligo-astheno-teratozoospermia restored by in vivo injection and electroporation of naked mRNA

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

The authors assess the effectiveness of electroporating mRNA into male germ cells to rescue the expression of proteins required for spermatogenesis progression in individuals where these proteins are mutated or depleted. To set up the methodology, they first evaluated the expression of reporter proteins in wild-type mice, which showed expression in germ cells for over two weeks. Then, they attempted to recover fertility in a model of late spermatogenesis arrest that produces immotile sperm. By electroporating the mutated protein, the authors recovered the motility of ~5% of the sperm; although the sperm regenerated was not able to produce offspring using IVF, the embryos reached the 2-cell state (in contrast to controls that did not progress past the zygote state).

This is a comprehensive evaluation of the mRNA methodology with multiple strengths. First, the authors show that naked synthetic RNA, purchased from a commercial source or generated in the laboratory with simple methods, is enough to express exogenous proteins in testicular germ cells. The authors compared RNA to DNA electroporation and found that germ cells are efficiently electroporated with RNA, but not DNA. The differences between these constructs were evaluated using in vivo imaging to track the reporter signal in individual animals through time. To understand how the reporter proteins affect the results of the experiments, the authors used different reporters: two fluorescent (eGFP and mCherry) and one bioluminescent (Luciferase). Although they observed differences among reporters, in every case expression lasted for at least two weeks. The authors used a relevant system to study the therapeutic potential of RNA electroporation. The ARMC2-deficient animals have impaired sperm motility phenotype that affects only the later stages of spermatogenesis. The authors showed that sperm motility was recovered to ~5%, which is remarkable due to the small fraction of germ cells electroporated with RNA with the current protocol. The sperm motility parameters were thoroughly assessed by CASA. The 3D reconstruction of an electroporated testis using state-of-the-art methods to show the electroporated regions is compelling.

The main weakness of the manuscript is that although the authors manage to recover motility in a small fraction of the sperm population, it is unclear whether the increased sperm quality is substantial to improve assisted reproduction outcomes. The authors found that the rescued sperm could be used to obtain 2-cell embryos via IVF, but no evidence for more advanced stages of embryo differentiation was provided. The motile rescued sperm was also successfully used to generate blastocyst by ICSI, but the statistical significance of the rate of blastocyst production compared to non-rescued sperm remains unclear. The title is thus an overstatement since fertility was never restored for IVF, and the mutant sperm was already able to produce blastocysts without the electroporation intervention.

Overall, the authors clearly show that electroporating mRNA can improve spermatogenesis as demonstrated by the generation of motile sperm in the ARMC2 KO mouse model.

We thank the reviewer for this thoughtful and constructive comment. We agree that our study demonstrates a partial functional recovery of spermatogenesis rather than a complete restoration of fertility. Our main objective was to establish and validate a proof-of-concept approach showing that mRNA electroporation can rescue the expression of a missing or mutated protein in post-meiotic germ cells and result in the production of motile sperm.

To address the reviewer’s concern, we have the title and discussion to more accurately reflect the scope of our findings. The new title reads:

“Sperm motility in mice with oligo-astheno-teratozoospermia restored by in vivo injection and electroporation of naked mRNA”

In the manuscript, we now emphasize that while motility recovery was significant, complete fertility restoration was not achieved. We have also clarified that:

The 5% recovery in motile sperm represents a substantial improvement considering the small population of germ cells reached by the current electroporation method.

The 2-cell embryo formation observed after IVF serves as a strong indication of partial functional recovery

Finally, we now explicitly state in the Discussion that this approach should be considered a therapeutic proof-of-concept, demonstrating feasibility and potential, rather than a fully curative intervention.

Reviewer #2 (Public review):

The authors inject, into the rete testes, mRNA and plasmids encoding mRNAs for GFP and then ARMC2 (into infertile Armc2 KO mice) in a gene therapy approach to express exogenous proteins in male germ cells. They do show GFP epifluorescence and ARMC2 protein in KO tissues, although the evidence presented is weak. Overall, the data do not necessarily make sense given the biology of spermatogenesis and more rigorous testing of this model is required to fully support the conclusions, that gene therapy can be used to rescue male infertility.

In this revision, the authors attempt to respond to the critiques from the first round of reviews. While they did address many of the minor concerns, there are still a number to be addressed. With that said, the data still do not support the conclusions of the manuscript.

We thank the reviewer for their careful and detailed assessment of our manuscript. We appreciate the concerns raised regarding mRNA stability, GFP localization, and the interpretation of spermatogenesis stages, and we have addressed these points in the manuscript and in the responses below.

(1) The authors have not satisfactorily provided an explanation for how a naked mRNA can persist and direct expression of GFP or luciferase for ~3 weeks. The most stable mRNAs in mammalian cells have half-lives of ~24-60 hours. The stability of the injected mRNAs should be evaluated and reported using cell lines. GFP protein's half-life is ~26 hours, and luciferase protein's half-life is ~2 hours.

We thank the reviewer for this important comment. The stability of mRNA-GFP was assessed by RT-QPCR in HEK cells and seminiferous tubule cells (Fig. 5). mRNA-GFP was detected for up to 60 hours in HEK cells and for up to two weeks in seminiferous tubule cells (Fig. 5A). Together, these results suggest that the long-lasting fluorescence observed in our experiments reflects a combination of transcript stability, efficient translation within germ cells and the slow protein turnover that is typical of the spermatogenic lineage.

(2) There is no convincing data shown in Figs. 1-8 that the GFP is even expressed in germ cells, which is obviously a prerequisite for the Armc2 KO rescue experiment shown in the later figures! In fact, to this reviewer the GFP appears to be in Sertoli cell cytoplasm, which spans the epithelium and surrounds germ cells - thus, it can be oft-confused with germ cells. In addition, if it is in germ cells, then the authors should be able to show, on subsequent days, that it is present in clones of germ cells that are maturing. Due to intracellular bridges, a molecule like GFP has been shown to diffuse readily and rapidly (in a matter of minutes) between adjacent germ cells. To clarify, the authors must generate single cell suspensions and immunostain for GFP using any of a number of excellent commercially-available antibodies to verify it is present in germ cells. It should also be present in sperm, if it is indeed in the germline.

We thank the reviewer for this insightful comment. To directly address the concern, we performed additional experiments to assess GFP expression in germ cells following in vivo mRNA delivery. GFP-encoding mRNA was injected and electroporated into the testes on day 0. On day 1, testes were collected, enzymatically dissociated, and the resulting seminiferous tubule cell suspensions were cultured for 12 hours. Live cells were then analyzed by fluorescence microscopy (Fig. 10).

We observed GFP expression in various germ cell types, including pachytene spermatocytes (53,4 %) (Fig 10 A-), round spermatids (25 %) (Fig 10B-E) and in elongated spermatids (11,4%) (Fig 10 C-E). The identification of these cell types was based on DAPI nuclear staining patterns, cell size fig 10 F, non-adherent characteristics, and the use of an enzymatic dissociation protocol.

Fluorescence imaging revealed strong cytoplasmic GFP signals in each of these populations, confirming efficient transfection and translation of the delivered mRNA. These results demonstrate that the in vivo injection and electroporation protocol enables effective mRNA transfection across multiple stages of spermatogenesis. These results confirm that the injected mRNA is efficiently translated in germ cells at various stages of spermatogenesis. Together, these data validate the germ cell-specific nature of the GFP signal, supporting the Armc2 KO rescue experiments.

As mentioned previously, we assessed the stability of mRNA-GFP using RT-QPCR in HEK cells and seminiferous tubule cells (see Fig. 5). mRNA-GFP was detected for up to 60 hours in HEK cells and for up to two weeks in seminiferous tubule cells. Together, these results suggest that the long-lasting fluorescence observed in our experiments reflects a combination of transcript stability and local translation within germ cells, as well as the slow protein turnover typical of the spermatogenic lineage.

Other comments:

70-1 This is an incorrect interpretation of the findings from Ref 5 - that review stated there were ~2,000 testis-enriched genes, but that does not mean "the whole process involves around two thousand of genes"

We thank the reviewer for this helpful comment. We agree that our previous phrasing was imprecise. We have revised the sentence to clarify that approximately 2,000 genes show testis-enriched expression, rather than implying that the entire spermatogenic process is limited to these genes. The corrected sentence now reads:

“Spermatogenesis involves the coordinated expression of a large number of genes, with approximately 2,000 showing testis-enriched expression, about 60% of which are expressed exclusively in the testes”

74 would specify 'male':

we have now specified it as you suggested.

79-84 Are the concerns with ICSI due to the procedure itself, or the fact that it's often used when there is likely to be a genetic issue with the male whose sperm was used? This should be clarified if possible, using references from the literature, as this reviewer imagines this could be a rather contentious issue with clinicians who routinely use this procedure, even in cases where IVF would very likely have worked:

We thank the reviewer for this important comment. Concerns about ICSI outcomes indeed reflect two partly overlapping causes: the procedure itself (direct sperm injection and associated laboratory manipulations) and the clinical/genetic background of couples undergoing ICSI (especially men with severe male-factor infertility). Large reviews and meta-analyses report a small increase in some perinatal and congenital risks after ART/ICSI, but these studies conclude that it is difficult to fully disentangle procedural effects from parental factors. Importantly, genetic or epigenetic abnormalities in the male (which motivate use of ICSI) likely contribute to adverse outcomes in offspring, while some studies also suggest that ICSI-specific manipulations may alter epigenetic marks in embryos. For these reasons professional bodies recommend reserving ICSI for appropriate male-factor indications rather than as routine insemination for non-male-factor cases

We have revised the text accordingly to clarify this distinction:

“ICSI can efficiently overcome the problems faced. Nevertheless, concerns persist regarding the potential risks associated with this technique, including blastogenesis defect, cardiovascular defect, gastrointestinal defect, musculoskeletal defect, orofacial defect, leukemia, central nervous system tumors, and solid tumors [1-4]. Statistical analyses of birth records have demonstrated an elevated risk of birth defects, with a 30-40 % increased likelihood in cases involving ICSI [1-4], and a prevalence of birth defects between 1 % and 4 % [3]. It is important to note, however, that the origin of these risks remains debated. Several large epidemiological and mechanistic studies indicate that both the procedure itself (direct microinjection and in vitro manipulation) and the underlying genetic or epigenetic abnormalities often present in men requiring ICSI contribute to the observed outcomes [1, 3] [5, 6] . To overcome these drawbacks, a number of experimental strategies have been proposed to bypass ARTs and restore spermatogenesis and fertility, including gene therapy [7-10].”

199 Codon optimization improvement of mRNA stability needs a reference;

We have added the references accordingly: [11-15]

In one study using yeast transcripts, optimization improved RNA stability on the order of minutes (e.g., from ~5 minutes to ~17 minutes); is there some evidence that it could be increased dramatically to days or weeks?

We agree with the reviewer that codon optimization can enhance mRNA stability, but available evidence indicates that this effect is moderate. In Saccharomyces cerevisiae, Presnyak et al. (2015) [16] showed that codon optimization increased mRNA half-life from approximately 5 minutes to ~17 minutes, representing a several-fold improvement rather than a shift to days or weeks. Similar codon-dependent stabilization has been observed in mammalian systems, where transcripts enriched in optimal codons exhibit longer half-lives and enhanced translation efficiency [11]; [17]). However, these studies consistently report effects on the scale of minutes to hours. In mammalian cells, the prolonged stability of therapeutic or vaccine mRNAs—lasting for days—is primarily achieved through additional features such as optimized untranslated regions, chemical nucleotide modifications (e.g., N¹-methylpseudouridine), and protective delivery systems, rather than codon usage alone ([18]; [19]).

Other molecular optimizations that improve in vivo mRNA stability and translation include a poly(A) tail, which binds poly(A)-binding proteins to protect the transcript from 3′ exonuclease degradation and promotes ribosome recycling, and a CleanCap structure at the 5′ end, which mimics the natural Cap 1 configuration, protects against 5′ exonuclease attack, and enhances translational initiation [11-15]. Together, these modifications act synergistically to stabilize the transcript and support efficient translation.

472-3 The reported half-life of EGFP is ~36 hours - so, if the mRNA is unstable (and not measured, but certainly could be estimated by qRT-PCR detection of the transcript on subsequent days after injection) and EGFP is comparatively more stable (but still hours), how does EGFP persist for 21 days after injection of naked mRNA??

We thank the reviewer for this important comment. The stability of mRNA-GFP was assessed by RT-QPCR in HEK cells and seminiferous tubule cells (Fig. 5). mRNA-GFP was detected for up to 60 hours in HEK cells and for up to two weeks in seminiferous tubule cells (Fig. 5). Together, these results suggest that the long-lasting fluorescence observed in our experiments reflects a combination of transcript stability, efficient translation within germ cells and the slow protein turnover that is typical of the spermatogenic lineage.

Curious why the authors were unable to get anti-GFP to work in immunostaining?

We appreciate the reviewer’s question. We attempted to detect GFP using several commercially available anti-GFP antibodies under various standard immunostaining conditions. However, in our hands, these antibodies consistently produced either no signal or high background staining, making the results unreliable. We therefore relied on direct detection of GFP fluorescence, which provides a more accurate and specific readout of protein expression in our system.

In Fig. 3-4, the GFP signals are unremarkable, in that they cannot be fairly attributed to any structure or cell type - they just look like blobs; and why, in Fig. 4D-E, why does the GFP signal appear stronger at 21 days than 15 days? And why is it completely gone by 28 days? This data is unconvincing.

We would like to thank the reviewer for their comments. Figure 3–4 provides a global overview of GFP expression on the surface of the testis. The entire testis was imaged using an inverted epifluorescence microscope, and the GFP signal represents a composite of multiple seminiferous tubules across the tissue surface. Due to this whole-organ imaging approach, it is not possible to resolve individual structures such as the basement membrane or lumen, which is why the signals may appear as diffuse “blobs.”

Regarding the time-course in Figure 4D–E, the apparent increase in GFP signal at 21 days compared with 15 days likely reflects accumulation and translation of the delivered mRNA in germ cells over time, whereas the absence of signal at 28 days corresponds to the natural turnover and degradation of GFP protein and mRNA in the tissue. We hope this explanation clarifies the observed patterns of fluorescence.

If the authors did a single cell suspension, what types or percentage of cells would be GFP+? Since germ cells are not adherent in culture, a simple experiment could be done whereby a single cell suspension could be made, cultured for 4-6 hours, and non-adherent cells "shaken off" and imaged vs adherent cells. Cells could also be fixed and immunostained for GFP, which has worked in many other labs using anti-GFP.

We thank the reviewer for this insightful comment. To directly address the concern, we performed additional experiments to assess GFP expression in germ cells following in vivo mRNA delivery. GFP-encoding mRNA was injected and electroporated into the testes on day 0. On day 1, testes were collected, enzymatically dissociated, and the resulting seminiferous tubule cell suspensions were cultured for 12 hours. Live cells were then analyzed by fluorescence microscopy (Fig. 10).

We observed GFP expression in various germ cell types, including pachytene spermatocytes (53,4 %) (Fig 10 A-), round spermatids (25 %) (Fig 10B-E) and in elongated spermatids (11,4%) (Fig 10 C-E). The identification of these cell types was based on DAPI nuclear staining patterns, cell size fig 10 F, non-adherent characteristics, and the use of an enzymatic dissociation protocol.

Fluorescence imaging revealed strong cytoplasmic GFP signals in each of these populations, confirming efficient transfection and translation of the delivered mRNA. These results demonstrate that the in vivo injection and electroporation protocol enables effective mRNA transfection across multiple stages of spermatogenesis.

These results confirm that the injected mRNA is efficiently translated in germ cells at various stages of spermatogenesis. Together, these data validate the germ cell-specific nature of the GFP signal, supporting the Armc2 KO rescue experiments.

As mentioned previously, we assessed the stability of mRNA-GFP using RT-QPCR in HEK cells and seminiferous tubule cells (see Fig. 5). mRNA-GFP was detected for up to 60 hours in HEK cells and for up to two weeks in seminiferous tubule cells. Together, these results suggest that the long-lasting fluorescence observed in our experiments reflects a combination of transcript stability and local translation within germ cells, as well as the slow protein turnover typical of the spermatogenic lineage.

In Fig. 5, what is the half-life of luciferase? From this reviewer's search of the literature, it appears to be ~2-3 h in mammalian cells. With this said, how do the authors envision detectable protein for up to 20 days from a naked mRNA? The stability of the injected mRNAs should be shown in a mammalian cell line - perhaps this mRNA has an incredibly long half-life, which might help explain these results. However, even the most stable endogenous mRNAs (e.g., globin) are ~24-60 hrs.

We did not directly assess the stability of luciferase mRNA, but we evaluated the persistence of GFP mRNA, which was synthesized and optimized using the same sequence optimization and chemical modification strategy as the luciferase mRNA. In these experiments, mRNA-GFP was detectable in seminiferous tubule cells for up to two weeks after injection. We therefore expect a similar stability profile for the luciferase mRNA. These findings suggest that the prolonged fluorescence or bioluminescence observed in our study likely reflects a combination of factors, including enhanced transcript stability, local translation within germ cells, and the inherently slow protein turnover characteristic of the spermatogenic lineage.

527-8 The Sertoli cell cytoplasm is not just present along the basement membrane as stated, but also projects all the way to the lumina:

we clarified the sentence " Sertoli cells have an oval to elongated nucleus and the cytoplasm presents a complex shape (“tombstone” pattern) along the basement membrane, with long projections that extend toward the lumen."

529-30 This is incorrect, as round spermatids are never "localized between the spermatocytes and elongated spermatids" - if elongated spermatids are present, rounds are not - they are never coincident in the same testis section:

We thank the reviewer for this important comment and for drawing attention to the detailed staging of the seminiferous epithelium. We agree that the spatial organization of germ cells varies depending on the stage of spermatogenesis. While round spermatids (steps 1–8) and elongated spermatids (steps 9–16) are typically associated with distinct stages, transitional stages of the seminiferous epithelium can contain both cell types in close proximity, reflecting the continuous and overlapping nature of spermatid differentiation (Meistrich, 2013, Methods Mol. Biol. 927:299–307). We have revised the text to clarify this point, indicating that the relative positioning of germ cell types depends on the stage of the seminiferous cycle rather than implying their constant coexistence within the same tubule section.

Fig. 7. To this reviewer, all of the GFP appears to be in Sertoli cell cytoplasm In Figs 1-8 there is no convincing evidence presented that GFP is expressed in germ cells! In fact, it appears to be in Sertoli cells.

We thank the reviewer for their observation. As previously mentioned, we have included an additional experiment specifically demonstrating GFP expression in germ cells (fig 10). This new data provides clear evidence that the GFP signal is not restricted to Sertoli cells and confirms successful uptake and translation of GFP mRNA in germ cells.

Fig. 9 - alpha-tubuline?

We corrected the figure.

Fig. 11 - how was sperm morphology/motility not rescued on "days 3, 6, 10, 15, or 28 after surgery", but it was in some at 21 and 35? How does this make sense, given the known kinetics of male germ cell development??

We note the reviewer’s concern regarding the timing of motile sperm appearance. Variability among treated mice is expected because transfection efficiency differed between spermatogonia and spermatids. Full spermiogenesis requires ~15 days, and epididymal transit adds ~8 days, consistent with motile sperm appearing around 21 days post-injection in some mice.

And at least one of the sperm in the KO in Fig. B5 looks relatively normal, and the flagellum may be out-of-focus in the image? With only a few sperm for reviewers to see, how can we know these represent the population?

We thank the reviewer for their comment. Upon closer examination of the image, the flagellum of the spermatozoon in question is clearly abnormally short and this is not due to being out of focus. Furthermore, the supplementary figure shows that the KO consistently lacks normal spermatozoa. These defects are consistent with previous findings from our laboratory [22], confirming that the observed phenotype is representative of the KO population rather than an isolated occurrence.

Reviewer #3 (Public review):

Summary:

The authors used a novel technique to treat male infertility. In a proof-of-concept study, the authors were able to rescue the phenotype of a knockout mouse model with immotile sperm using this technique. This could also be a promising treatment option for infertile men.

Strengths:

In their proof-of-concept study, the authors were able to show that the novel technique rescues the infertility phenotype of Armc2 knockout spermatozoa. In the current version of the manuscript, the authors have added data on in vitro fertilisation experiments with Armc2 mRNA-rescued sperm. The authors show that Armc2 mRNA-rescued sperm can successfully fertilise oocytes that develop to the blastocyst stage. This adds another level of reliability to the data.

Weaknesses:

Some minor weaknesses identified in my previous report have already been fixed. The technique is new and may not yet be fully established for all issues. Nevertheless, the data presented in this manuscript opens the way for several approaches to immotile spermatozoa to ensure successful fertilisation of oocytes and subsequent appropriate embryo development.

[Editors' note: The images in Figure 12 do not support the authors' interpretation that 2-cell embryos resulted from in vitro fertilization. Instead, the cells shown appear to be fragmented, unfertilized eggs. Combined with the lack of further development, it seems highly unlikely that fertilization was successful.]

We thank the reviewer for their careful evaluation and constructive feedback. We appreciate the acknowledgment of the strengths of our study, particularly the proof-of-concept demonstration that Armc2-mRNA electroporation can rescue sperm motility in Armc2 knockout mice.

Regarding the concern raised by the editor about Figure 12, we would like to clarify two technical points. First, the IVF experiments were performed using CD1 oocytes and B6D2 sperm. Due to strain-specific incompatibilities, fertilization of CD1 oocytes by B6D2 sperm typically does not progress beyond the two-cell stage (Fernández-González [23] et al., 2008, Biology of Reproduction). Therefore, the observation of two-cell embryos represents the expected limit of development in this cross and serves as a strong indication of successful fertilization, even though further development is not possible. Second, the oocytes used in these experiments were treated with collagenase to remove cumulus cells. This enzymatic treatment can sometimes affect the morphology of early embryos, which may explain why the two-cell embryos in Figure 12 appear less regular or somewhat fragmented. We also included a control showing embryos from B6D2 sperm with the same collagenase treatment on CD1 oocytes, which yielded similar appearances (Fig14 A4).

To provide additional functional evidence, we complemented the IVF experiments with ICSI using rescued Armc2–/– sperm and B6D2 oocytes, which allowed embryos to develop to the blastocyst stage. In these experiments, 25% of injected oocytes reached the blastocyst stage with rescued sperm compared to 13% for untreated Armc2–/– sperm (Supplementary Fig. 9) These results support the functional competence of rescued sperm and demonstrate partial recovery of fertilization ability following Armc2 mRNA electroporation.

We have clarified these points in the revised Results and Discussion sections to emphasize that the IVF data indicate partial functional recovery of rescued sperm rather than full fertility restoration. These clarifications address the editor’s concern while accurately representing the technical limitations of the strain combination used in our experiments.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Fig 12 and Supplementary Fig 9 are mislabeled in the text and rebuttal.

We thank the reviewer for pointing this out. We have carefully checked the manuscript and the rebuttal text, and corrected all references to Figure 12 and Supplementary Figure 9 to ensure they are accurately labeled and consistent throughout the text.

Reviewer #3 (Recommendations for the authors):

The contribution of the newly added authors should be clarified. All other aspects of inadequacy raised in my previous report have been adequately addressed.

No further comments.

We thank the reviewer for noting this. The contributions of the newly added authors have been clarified in the Author Contributions section of the revised manuscript. All other points raised in the previous review have been addressed as indicated.

References

(1) Hansen, M., et al., Assisted reproductive technologies and the risk of birth defects--a systematic review. Hum Reprod, 2005. 20(2): p. 328-38.

(2) Halliday, J.L., et al., Increased risk of blastogenesis birth defects, arising in the first 4 weeks of pregnancy, after assisted reproductive technologies. Hum Reprod, 2010. 25(1): p. 59-65.

(3) Davies, M.J., et al., Reproductive technologies and the risk of birth defects. N Engl J Med, 2012. 366(19): p. 1803-13.

(4) Kurinczuk, J.J., M. Hansen, and C. Bower, The risk of birth defects in children born after assisted reproductive technologies. Curr Opin Obstet Gynecol, 2004. 16(3): p. 201-9.

(5) Graham, M.E., et al., Assisted reproductive technology: Short- and long-term outcomes. Dev Med Child Neurol, 2023. 65(1): p. 38-49.

(6) Palermo, G.D., et al., Intracytoplasmic sperm injection: state of the art in humans. Reproduction, 2017. 154(6): p. F93-f110.

(7) Usmani, A., et al., A non-surgical approach for male germ cell mediated gene transmission through transgenesis. Sci Rep, 2013. 3: p. 3430.

(8) Raina, A., et al., Testis mediated gene transfer: in vitro transfection in goat testis by electroporation. Gene, 2015. 554(1): p. 96-100.

(9) Michaelis, M., A. Sobczak, and J.M. Weitzel, In vivo microinjection and electroporation of mouse testis. J Vis Exp, 2014(90).

(10) Wang, L., et al., Testis electroporation coupled with autophagy inhibitor to treat non-obstructive azoospermia. Mol Ther Nucleic Acids, 2022. 30: p. 451-464.

(11) Wu, Q., et al., Translation affects mRNA stability in a codon-dependent manner in human cells. eLife, 2019. 8: p. e45396.

(12) Gallie, D.R., The cap and poly(A) tail function synergistically to regulate mRNA translational efficiency. Genes & Development, 1991. 5(11): p. 2108-2116.

(13) Henderson, J.M., et al., Cap 1 messenger RNA synthesis with co-transcriptional CleanCap® analog improves protein expression in mammalian cells. Nucleic Acids Research, 2021. 49(8): p. e42.

(14) Stepinski, J., et al., Synthesis and properties of mRNAs containing novel “anti-reverse” cap analogs. RNA, 2001. 7(10): p. 1486-1495.

(15) Sachs, A.B., P. Sarnow, and M.W. Hentze, Starting at the beginning, middle, and end: translation initiation in eukaryotes. Cell, 1997. 89(6): p. 831-838.

(16) Presnyak, V., et al., Codon optimality is a major determinant of mRNA stability. Cell, 2015. 160(6): p. 1111-24.

(17) Cao, D., et al., Unlock the sustained therapeutic efficacy of mRNA. J Control Release, 2025. 383: p. 113837.

(18) Karikó, K., et al., Incorporation of pseudouridine into mRNA yields superior nonimmunogenic vector with increased translational capacity and biological stability. Mol Ther, 2008. 16(11): p. 1833-40.

(19) Pardi, N., et al., mRNA vaccines — a new era in vaccinology. Nature Reviews Drug Discovery, 2018. 17(4): p. 261-279.

(20) Meistrich, M.L. and R.A. Hess, Assessment of Spermatogenesis Through Staging of Seminiferous Tubules, in Spermatogenesis: Methods and Protocols, D.T. Carrell and K.I. Aston, Editors. 2013, Humana Press: Totowa, NJ. p. 299-307.

(21) Au - Mäkelä, J.-A., et al., JoVE, 2020(164): p. e61800.

(22) Coutton, C., et al., Bi-allelic Mutations in ARMC2 Lead to Severe Astheno-Teratozoospermia Due to Sperm Flagellum Malformations in Humans and Mice. Am J Hum Genet, 2019. 104(2): p. 331-340.

(23) Fernández-Gonzalez, R., et al., Long-term effects of mouse intracytoplasmic sperm injection with DNA-fragmented sperm on health and behavior of adult offspring. Biol Reprod, 2008. 78(4): p. 761-72.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Oligo-astheno-teratozoospermia (OAT), a recurent cause of male infertility, is the most frequent disorder of spermatogenesis with a probable genetic cause. Patients and mice bearing mutations in the ARMC2 gene have a decreased sperm concentration, and individual sperm show multiple morphological defects and a lack of motility – a canonical OAT phenotype. Intracytoplasmic sperm injection (ICSI) is required to treat such a condition but it is associated with a small increase in birth defects in comparison to pregnancies not involving assisted conception . Consequently, new targeted treatments are needed to restore fertility. Here, a combination of in vivo injection and electroporation of capped and poly-A-tailed naked mRNA is tested as a strategy to treat ARMC2-related infertility in mouse. mRNAs coding for several reporter genes are tested and the efficiency and the kinetic of expression are assessed using in vivo and in vitro 2D and 3D imaging experiments. We show that mRNA-coded reporter proteins are detected for up to 3 weeks in germ cells, making the use of mRNA possible to treat infertility. We compare these results with those obtained with a non-integrative plasmid Enhanced Episomal Vector (EEV), which induces low and transient expression in spermatogenic cells. Consequently, injection and electroporation of naked mRNA-Armc2 into the testes of Armc2-deficient males were performed and we show the presence of normal and motile sperm in the epididymis. These motile sperm were able to produce embryos by IVF and ICSI. This study shows for the first time that mRNA-Armc2 efficiently restores fertility and opens new paths for male infertility treatment.

Revised version: Oligo-astheno-teratozoospermia (OAT), a recurent cause of male infertility, is the most frequent disorder of spermatogenesis with a probable genetic cause. Patients and mice bearing mutations in the ARMC2 gene have a decreased sperm concentration, and individual sperm show multiple morphological defects and a lack of motility – a canonical OAT phenotype. Intracytoplasmic sperm injection (ICSI) is required to treat such a condition but it is associated with a small increase in birth defects in comparison to pregnancies not involving assisted conception. Consequently, new targeted treatments are needed to restore fertility. Here, a combination of in vivo injection and electroporation of capped and poly-A-tailed naked mRNA is tested as a strategy to treat ARMC2-related infertility in mouse. mRNAs coding for several reporter genes are tested and the efficiency and the kinetic of expression are assessed using in vivo and in vitro 2D and 3D imaging experiments. We show that mRNA-coded reporter proteins are detected for up to 3 weeks in germ cells, making the use of mRNA possible to treat infertility. We compare these results with those obtained with a non-integrative plasmid Enhanced Episomal Vector (EEV), which induces low and transient expression in spermatogenic cells. Consequently, injection and electroporation of naked mRNA-Armc2 into the testes of Armc2-deficient males were performed and we show the presence of normal and motile sperm in the epididymis. These motile sperm were able to produce embryos by IVF and ICSI. This study demonstrates, for the first time, that mRNA electroporation can restore sperm motility and partially fertilizing ability, providing a proof-of-concept for mRNA-based strategies to correct monogenic causes of male infertility and opening new avenues for male infertility treatment.

Introduction

Worldwide, 10-15 % of couples (or 70 million) face infertility [1]. Infertility is thus a major public health issue presenting significant medical, scientific and economic challenges (a multibillion € annual market)[2]. A significant proportion of infertilities is due to altered gametogenesis, where the sperm and eggs produced are incompatible with fertilization and/or embryonic development [3]. Approximately 40 % of cases of infertilities involve a male factor, either exclusively, or associated with a female deficiency [4].

Earlier version: Male gametogenesis, or spermatogenesis, is a highly complex physiological process which can be split into three successive steps: proliferation (mitosis of spermatogonia), reduction of the number of chromosomes (meiosis of spermatocytes), and morphogenesis of spermatozoa (spermiogenesis). The whole process involves around two thousands of genes, 60 % of which are expressed exclusively in the testes [5]. Because of this multiplicity of genes, spermatogenesis is strongly affected by genetic factors [5], with most severe disorders likely to be of genetic origin.

Revised version: Male gametogenesis, or spermatogenesis, is a highly complex physiological process which can be split into three successive steps: proliferation (mitosis of spermatogonia), reduction of the number of chromosomes (meiosis of spermatocytes), and morphogenesis of spermatozoa (spermiogenesis). Spermatogenesis involves the coordinated expression of a large number of genes, with approximately 2,000 showing testis-enriched expression, about 60% of which are expressed exclusively in the testes [5]. Because of this multiplicity of genes, spermatogenesis is strongly affected by genetic factors [5], with most severe disorders likely to be of genetic origin.

Earlier version: Among infertility disorders, oligo-astheno-teratozoospermia (OAT) is the most frequent (50 % [6]) and it is likely to be of genetic origin. Spermatocytograms of OAT patients show a decrease in sperm concentration, multiple morphological defects and defective motility [7, 8]. Because of these combined defects, patients are infertile and can only conceive by intracytoplasmic sperm injection (ICSI).

Revised version: Among male infertility disorders, oligo-astheno-teratozoospermia (OAT) is the most frequent (50 % [6]) and it is likely to be of genetic origin. Spermatocytograms of OAT patients show a decrease in sperm concentration, multiple morphological defects and defective motility [7, 8]. Because of these combined defects, patients are infertile and can only conceive by intracytoplasmic sperm injection (ICSI).

Earlier version: ICSI can efficiently overcome the problems faced. Nevertheless, concerns persist regarding the potential risks associated with this technique, including blastogenesis defect, cardiovascular defect, gastrointestinal defect, musculoskeletal defect, orofacial defect, leukemia, central nervous system tumors, and solid tumors [9–12]. Statistical analyses of birth records have demonstrated an elevated risk of birth defects, with a 30-40 % increased likelihood in cases involving ICSI [9–12], and a prevalence of birth defects between 1 % and 4 % [11]. To overcome these drawbacks, a number of experimental strategies have been proposed to bypass ARTs and restore spermatogenesis and fertility, including gene therapy [13–16].

Revised version: ICSI can efficiently overcome the problems faced. Nevertheless, concerns persist regarding the potential risks associated with this technique, including blastogenesis defect, cardiovascular defect, gastrointestinal defect, musculoskeletal defect, orofacial defect, leukemia, central nervous system tumors, and solid tumors [9–12]. Statistical analyses of birth records have demonstrated an elevated risk of birth defects, with a 30-40 % increased likelihood in cases involving ICSI [9–12], and a prevalence of birth defects between 1 % and 4 % [11]. It is important to note, however, that the origin of these risks remains debated. Several large epidemiological and mechanistic studies indicate that both the procedure itself (direct microinjection and in vitro manipulation) and the underlying genetic or epigenetic abnormalities often present in men requiring ICSI contribute to the observed outcomes [9, 11] [13, 14]. To overcome these drawbacks, a number of experimental strategies have been proposed to bypass ARTs and restore spermatogenesis and fertility, including gene therapy [15–18].

Earlier version: Gene therapy consists of introducing a DNA sequence into the genome to compensate for a defective gene. It can thus rescue production of a missing protein, and is now applied both in research [17] and for the treatment of human diseases [18] .

Revised version: Gene therapy consists of introducing a DNA sequence into the genome to compensate for a defective gene. It can thus rescue production of a missing protein, and is now applied both in research [19] and for the treatment of human diseases [20].

Earlier version: Given the genetic basis of male infertility, the first strategy, tested in mice, was to overcome spermatogenic failure associated with monogenic diseases by delivery of an intact gene to deficient germ cells [13]. Gene therapy is effective in germ cells, as numerous publications have shown that conventional plasmids can be transferred into spermatogonia in several species with success, allowing their transcription in all cells of the germinal lineage [13–16]. Most experiments were performed in mouse models, delivering DNA constructs into living mouse germ cells by testis electroporation after microinjection of a DNA-containing solution into the seminiferous tubules. Using this method, it was possible to rescue meiosis and fertility in mouse models of infertility [13, 16]. However, the genetic changes induced are transmitted to any descendants. Consequently, gene therapy cannot be used to treat infertility in humans, both for ethical reasons and to avoid any eugenic deviations, and currently transmissible changes in humans are illegal in 39 countries [19]. Furthermore, the genetic modification of germ cell lines poses biological risks, including the induction of cancer, off-target effects, and cell mosaicism. Errors in editing may have adverse effects on future generations. It is exceedingly challenging to anticipate the consequences of genetic mosaicism, for instance, in a single individual [20, 21]. Gene therapies have thus raised both ethical controversy and long- term safety issues.

Revised version: Given the genetic basis of male infertility, the first strategy, tested in mice, was to overcome spermatogenic failure associated with monogenic diseases by delivery of an intact gene to deficient germ cells [15]. Gene therapy is effective in germ cells, as numerous publications have shown that conventional plasmids can be transferred into spermatogonia in several species with success, allowing their transcription in all cells of the germinal lineage [15–18]. Most experiments were performed in mouse models, delivering DNA constructs into living mouse germ cells by testis electroporation after microinjection of a DNA-containing solution into the seminiferous tubules. Using this method, it was possible to rescue meiosis and fertility in mouse models of infertility [15, 18]. However, the genetic changes induced are transmitted to any descendants. Consequently, gene therapy cannot be used to treat infertility in humans, both for ethical reasons and to avoid any eugenic deviations, and currently transmissible changes in humans are illegal in 39 countries [21]. Furthermore, the genetic modification of germ cell lines poses biological risks, including the induction of cancer, off-target effects, and cell mosaicism. Errors in editing may have adverse effects on future generations. It is exceedingly challenging to anticipate the consequences of genetic mosaicism, for instance, in a single individual [22, 23]. Gene therapies have thus raised both ethical controversy and long-term safety issues.

Earlier version: For these reasons, we decided to test an alternative strategy to DNA transfection based on the use of naked mRNA. Thanks to this change, the risk of genomic insertion is avoided, and thus there is no question of heritable alterations [22].The first part of this study presents a characterization of the protein expression patterns obtained following transfection of naked mRNA coding for reporter genes into the testes of mice. The second part is to apply the protocol to a preclinical mouse model of OAT. Patients and mice carrying mutations in the ARMC2 gene present a canonical OAT phenotype and are infertile. The preclinical Armc2 deficient (Armc2 KO) mouse model is therefore a valuable model to assess whether in vivo injection of naked mRNA combined with electroporation can restore spermatogenesis. We chose this model for several reasons: first, Armc2 KO mice are sterile and all sperm exhibit short, thick or coiled flagella [13]. As a result, 100 % of sperm are immobile, thus it should be easy to determine the efficacy of the technique by measuring sperm motility with a CASA system. Second, the Armc2 gene codes for an 867-amino acid protein and this large size represents a challenge for expression in the testis following electroporation.

Revised version: For these reasons, we decided to test an alternative strategy to DNA transfection based on the use of naked mRNA. Thanks to this change, the risk of genomic insertion is avoided, and thus there is no question of heritable alterations [24].The first part of this study presents a characterization of the protein expression patterns obtained following transfection of naked mRNA coding for reporter genes into the testes of mice. The second part is to apply the protocol to a preclinical mouse model of OAT. Patients and mice carrying mutations in the ARMC2 gene present a canonical OAT phenotype and are infertile. The preclinical Armc2 deficient (Armc2 KO) mouse model is therefore a valuable model to assess whether in vivo injection of naked mRNA combined with electroporation can restore spermatogenesis. We chose this model for several reasons: first, Armc2 KO mice are sterile and all sperm exhibit short, thick or coiled flagella [13]. As a result, 100 % of sperm are immobile, thus it should be easy to determine the efficacy of the technique by measuring sperm motility with a CASA system. Second, the Armc2 gene codes for an 867-amino acid protein and this large size represents a challenge for expression in the testis following electroporation.

To determine the efficacy of naked mRNA transfection as a method to achieve protein expression in the testes, we first assessed the level of transcription of reporter proteins after mRNA injection compared to the injection of a non-integrating plasmid, the Enhanced Episomal Vector (EEV). EEV is a vector derived from Epstein-Barr virus, it includes an origin of replication (EBV Ori) and the EBNA1 protein. Both elements, allows the synchronous initiation of extra-chromosomal EEV replication with host DNA at each S phase of the cell cycle and the segregation of the EEV episome in daughter cells. It is notable that EEV is maintained at a rate of 90-95% per cell division. It does not integrate or modify the host genome [11, 12].

In the present in vivo work, we injected and electroporated three distinct mRNAs coding for the following reporter proteins: green fluorescent protein (GFP), luciferase (Luc) and mCherry, and an EEV episome vector containing the sequences coding for both GFP and luciferase reporter proteins. The initial step was to characterize and validate the method of injection in adult males. Subsequently, the kinetics and patterns of expression of the electroporated mRNAs and EEV were compared using a variety of methods, including whole testis imaging, in vivo bioluminescence imaging, and tissue clearing. Subsequently, the mRNA transfection protocol was tested in a preclinical mouse model of OAT with the objective of restoring fertility.

Materials and methods

Animals

All procedures involving animals were performed in line with the French guidelines for the use of live animals in scientific investigations. The study protocol was approved by the local ethics committee (ComEth Grenoble # 318) and received governmental authorization (ministerial agreement # 38109-2022072716142778).

Earlier version: The animals used were (a) B6D2 F1 hybrid (♀ C57BL/6JRj crossed with ♂ DBA/2, Janvier laboratories; France) adult male mice aged between 8 and 25 weeks, (b) the Armc2 KO mouse strain obtained by CRISPR-Cas9 [23] and (c) CD1 female 6 weeks old (Janvier laboratories, Le Genest-Saint-Isle ,France). Experiments were carried out on wild type (WT) or Armc2 KO adult male mice aged between 8 and 15 weeks.

Revised version: The animals used were (a) B6D2 F1 hybrid (♀ C57BL/6JRj crossed with ♂ DBA/2, Janvier laboratories; France) adult male mice aged between 8 and 25 weeks, (b) the Armc2 KO mouse strain obtained by CRISPR-Cas9 [25] and (c) CD1 female 6 weeks old (Janvier laboratories, Le Genest-Saint-Isle, France). Experiments were carried out on wild type (WT) or Armc2 KO adult male mice aged between 8 and 15 weeks.

Chemicals and reagents

Earlier version: Fast Green (FG) (F7258 – 25 g), phosphate buffered saline (PBS, D853 7-500 mL), hematoxylin (GH5316 – 500 mL), eosin (HT110216 – 500 mL), tert-Butanol (471712 – 100 mL), Histodenz (D2158 – 100 g), sorbitol (S1876 – 100 g), urea (U5128 – 500 g), Potassium Phosphate, Monobasic (P0662), Magnesium Sulfate (anhydrous) (M9397), Sodium Bicarbonate Calcium Chloride ⋅ 2H2O (22317), EDTA (E9884), Sodium Lactate (60 % syrup - d= 1,32 g L-1) (L7900), Glucose (G8270), Penicillin (P4333), Streptomycin (P4333), HEPES (H0887), PVA 30,000-70,000 (P8136), Albumin, Bovine Fraction V (A3803), M2 medium (M7167), Hyaluronidase (H3884), mineral oil (M8410), PolyVinylPyrolidone (PVP360-100G), M16 medium (MR-016) and Collagenase (C8176) were purchased from Sigma Aldrich (Saint-Quentin-Fallavier, France). Sodium Chloride (1112-A) was purchased from Euromedex (Souffelweyersheim, France). Potassium Chloride (26764) L-Glutamine (35050038), Sodium Pyruvate (11360039), NEAA (11140050) and EAA (11130036) were purchased from Life Technologies, (Waltham, MA USA). Schorr staining solution was obtained from Merck (Darmstadt, Germany). Mfel HF (R35895) and RNase-free DNase I (M03035) were obtained from New England Biolabs (Ipswich, MA, USA). Paraformaldehyde (PFA, 15710) was obtained from Electron Microscopy Science (Hatfield, PA, USA). Ketamine and xylazine were obtained from Centravet (Dinan, France). Fluorescent i-particles (NIRFiP-180) were obtained from Adjuvatis (Lyon, France). Script Gurad RNAse CleanCap AG (040N-7113-1), CleanCAp EGFP-mRNA (040L-7601-100), T7-FlashScribe kit (C-ASF3507), poly(A) tail (C-PAP5104H) and CleanCap Luciferase-mRNA (L-7602-1000) were obtained from Tebubio (Le Perray en Yvelines, France).

Revised version: Fast Green (FG) (F7258 – 25 g), phosphate buffered saline (PBS, D853 7-500 mL), hematoxylin (GH5316 – 500 mL), eosin (HT110216 – 500 mL), tert-Butanol (471712 – 100 mL), Histodenz (D2158 – 100 g), sorbitol (S1876 – 100 g), urea (U5128 – 500 g), Potassium Phosphate, Monobasic (P0662), Magnesium Sulfate (anhydrous) (M9397), Sodium Bicarbonate Calcium Chloride ⋅ 2H2O (22317), EDTA (E9884), Sodium Lactate (60 % syrup - d= 1.32 g L−1) (L7900), Glucose (G8270), Penicillin (P4333), Streptomycin (P4333), HEPES (H0887), PVA 30,000-70,000 (P8136), Albumin, Bovine Fraction V (A3803), RPMI-1640 media (R2405), M2 medium (M7167), Hyaluronidase (H3884), mineral oil (M8410), PolyVinylPyrolidone (PVP360-100G), M16 medium (MR-016), Trypsin (T6567) and Collagenase (C8176) were purchased from Sigma Aldrich (Saint-Quentin-Fallavier, France). Sodium Chloride (1112-A) was purchased from Euromedex (Souffelweyersheim, France). Potassium Chloride (26764), L-Glutamine (35050038), Hoechst 33342 (R37165), Sodium Pyruvate (11360039), KSR serum (10828010), NEAA (11140050) and EAA (11130036) were purchased from Life Technologies, (Waltham, MA USA). Schorr staining solution was obtained from Merck (Darmstadt, Germany). Mfel HF (R35895) and RNase-free DNase I (M03035) were obtained from New England Biolabs (Ipswich, MA, USA). Paraformaldehyde (PFA, 15710) was obtained from Electron Microscopy Science (Hatfield, PA, USA). Ketamine and xylazine were obtained from Centravet (Dinan, France). Fluorescent i-particles (NIRFiP-180) were obtained from Adjuvatis (Lyon, France). Script Gurad RNAse CleanCap AG (040N-7113-1), CleanCAp EGFP-mRNA (040L-7601-100), T7-FlashScribe kit (C-ASF3507), poly(A) tail (C-PAP5104H) and CleanCap Luciferase-mRNA (L-7602-1000) were obtained from Tebubio (Le Perray en Yvelines, France).

Plasmids

All plasmids, EEV CAGs-GFP-T2A-Luciferase, (EEV604A-2; System Bioscience, Palo Alto, CA, USA), mCherry plasmid (given by Dr. Conti MD at UCSF, San Francisco, CA, USA) and EEV-Armc2-GFP plasmid (CUSTOM-S017188-R2-3, Trilink, San Diego, CA, USA) were amplified by bacterial transformation (E. coli, EC0112; Thermo Fisher, Courtaboeuf, France). After expansion, plasmids were purified with a NucleoBond Xtra Midi kit (740410-50; Macherey-Nagel, Düren, Germany) using manufacturer’s protocol. All plasmids DNA pellets were solubilized in (DNAse- and RNAse-free) milliQ water (Sigma Aldrich, Saint-Quentin-Fallavier, France), before being used.

The EEV CAGs-GFP-T2A-Luciferase episome contains the cDNA sequences of Green Fluorescent Protein (GFP) and luciferase, under the control of a CAGs promoter (Supp Fig 1). After purification, the EEV CAGs-GFP-T2A-Luciferase plasmid concentration was adjusted to 9 μg μL−1. Prior to injection, 3.3 μL of this plasmid solution was mixed with 1 μL 0.5 % Fast Green and 5.7 µL sterile PBS to obtain a final EEV concentration of 3 µg μL−1. The EEV-Armc2-GFP plasmid contains the mouse cDNA sequences of Armc2 (ENSMUST00000095729.11) and the Green Fluorescent Protein (GFP) genes under the control of a strong CAGs promoter (Supp Fig 1). After amplification and purification, the final plasmid concentration was adjusted to 9 μg μL−1 in water. Prior to injection, 3.3 μL of this plasmid solution was mixed with 1 μL of 0.5 % Fast Green and 5.7 µL of sterile PBS to obtain a final EEV concentration of 3 µg µL−1. The mCherry plasmid contains the cDNA sequence of mCherry under the control of CMV and T7 promoters (Supp Fig 1). After amplification and purification, the final plasmid concentration was adjusted to 9 μg μL−1.

Added in the revised version: Distribution of I-particles NIRFIP-180 in testis injected via the rete testis route.

Added in the revised version: (A) A solution containing 0.05 % Fast Green and 1 % fluorescent i-particles NIRFiP-180 was prepared, 10 µL was injected into the seminiferous tubules of adult males, through the rete testes and its efferent channels. Injection was performed at constant pressure under a binocular microscope. The progression of filling of the seminiferous tubules was monitored thanks to the Fast Green. (B) The testes were only filled to 2/3 capacity in order to prevent damage to the tissue. (C) Representative distribution of fluorescent i-particles NIRFiP-180 in a whole cross-section of an injected testis. Nuclei were counterstained with DAPI (blue emission) to reveal tubules. (D) Enlargement of a seminiferous tubule showing particles localized inside the lumens of the tubules. Scales bars: 1 mm and 500 µm.

Armc2-mRNA

Earlier version: Armc2-mRNA used for in vivo testes microinjection and electroporation was obtained from Trilink (San Diego, CA, USA). The commercial Armc2-mRNA has an AG CleanCap, a poly (A) tail of 120 adenosines and 3 HA-FLAG. The main challenge with mRNA-based therapy is mRNA stability. To improve mRNA stability in vivo and avoid its degradation by ribonucleases, optimization techniques were implemented. Thus, the used mRNA has codon optimization, a poly (A) tail and a CleanCap. To verify the efficiency of cell transfection, an EGFP-mRNA was injected together with the Armc2-mRNA. During the injection, the concentration of EGFP- mRNA and Armc2-mRNA were 300 ng µL-1 each.

Revised version: Armc2-mRNA used for in vivo testes microinjection and electroporation was obtained from Trilink (San Diego, CA, USA). The commercial Armc2-mRNA has an AG CleanCap, a poly (A) tail of 120 adenosines and 3 HA-FLAG. The main challenge with mRNA-based therapy is mRNA stability. To improve mRNA stability in vivo and avoid its degradation by ribonucleases, optimization techniques were implemented. Thus, the used mRNA has codon optimization, a poly (A) tail and a CleanCap[26–30]. To verify the efficiency of cell transfection, an EGFP-mRNA was injected together with the Armc2-mRNA. During the injection, the concentration of EGFP-mRNA and Armc2-mRNA were 300 ng µL−1 each.

mCherry-mRNA transcription in vitro

The circular mCherry plasmid was linearized using the restriction enzyme Mfel HF at 37 °C for 1 h. The pm-mCherry was then extracted and purified with the DNA Clean and Concentrator-25 TM kit (D4033; Zymo Research, Irvine, CA, USA). The pm-mCherry was subsequently transcribed in vitro using the T7-FlashScribe kit (C-ASF3507; Tebubio, Le Perray en Yvelines, France). The mRNA was capped with a clean cap (CleanCap AG; 040N-7113-1, Tebubio, Le Perray en Yvelines, France), and a poly(A) tail (C-PAP5104H; Tebubio, Le Perray en Yvelines, France) was added before purification using the NucleoSpin RNA Clean Up kit (740948-50; Macherey-Nagel, Düren, Germany). Prior to injection, mCherry-mRNA was mixed with Fast Green to obtain a final concentration of 300 ng μL−1 (0.05 % FG, PBS).

Agarose gel electrophoresis of the Episomal Vector EEV and mRNA-mCherry

Before loading on a pre-stained (Gel Green) 1.5 % agarose gel, the EEV-plasmid and mRNA were mixed with a loading dye (Promega Blue Orange G1904, Promega, Charbonnières France). An aliquot of each sample (500 ng) was loaded into each well and electrophoresis was performed for 30 min at 100 V at room temperature (RT). A DNA size marker (Gene ruler SM1331, Thermo Scientific, Courtaboeuf, France) was used to assess molecular weight. Gel images were acquired using a ChemiDoc XRS+ (BIORAD, Marnes-la-Coquette, France).

In vivo microinjection and electroporation of testes

Earlier version: Electroporation was conducted as previously described [15]. Briefly, male mice B6D2, Armc2+/+ or Armc2-/-, depending on the experimental conditions, were anesthetized with ketamine/xylazine solution (100 mg μL-1 and 10 mg μL-1 respectively). The testes were pulled out of the abdominal cavity or scrotum. Under a binocular microscope and using microcapillaries pipettes (FemtoTip II ®, Eppendorf, Montesson, France), 10 μL DNA (3 μg μL- 1 - 0.05 % FG) or mRNA (300 ng μL-1 -0.05 % FG) was injected into the seminiferous tubules through the rete testis applying a constant pressure (microinjector, Femto Jet 4i, Eppendorf, Montesson, France). Two series of 8 square electric pulses (25 V for 50 ms) were applied to the testis using tweezer-type electrodes linked to an electroporator (Gemini, BTX, Holliston, MA, USA). The testes were then replaced in the abdominal cavity, and the abdominal wall and skin were sutured. For each animal, the left testis was injected and electroporated with the different nucleic acids (mRNA, EEV), whereas the right testis was injected with a control solution (PBS, 0.5 % FG) as a control. Both testes were electroporated.

Revised version: Electroporation was conducted as previously described [17]. Briefly, male mice B6D2, Armc2+/+ or Armc2−/−, depending on the experimental conditions, were anesthetized with ketamine/xylazine solution (100 mg μL−1 and 10 mg μL−1 respectively). The testes were pulled out of the abdominal cavity or scrotum. Under a binocular microscope and using microcapillaries pipettes (FemtoTip II ®, Eppendorf, Montesson, France), 10 μL DNA (3 μg μL−1 - 0.05 % FG) or mRNA (300 ng μL−1 -0.05 % FG) was injected into the seminiferous tubules through the rete testis applying a constant pressure (microinjector, Femto Jet 4i, Eppendorf, Montesson, France). Two series of 8 square electric pulses (25 V for 50 ms) were applied to the testis using tweezer-type electrodes linked to an electroporator (Gemini, BTX, Holliston, MA, USA). The testes were then replaced in the abdominal cavity, and the abdominal wall and skin were sutured. For each animal, the left testis was injected and electroporated with the different nucleic acids (mRNA, EEV), whereas the right testis was injected with a control solution (PBS, 0.5 % FG) as a control. Both testes were electroporated.

Flash Freezing and Fluorescence Analysis of testes

Depending on the experimental condition, 1-, 3- or 7-days post injection, the testes were collected and washed for 5 min in PBS. Then, they were embedded in Peel-A-Way Cryomolds filled with OCT Mounting media (VWR, Rosny-sous-Bois, France). The samples were flash frozen in a 100 % isopentane solution (524391; Carlo ERBA, Val-de-Reuil, France), pre-cooled with liquid nitrogen. Once frozen, they were cut into 20 µm sections using a cryostat. The cryostat sections were then fixed with 4 % PFA-PBS for 10 min at 4 °C and counterstained with 1.8 μM DAPI (nuclear stain) before observation using an Axioscan Z1 slide scanner or a confocal microscope LSM710 (NLO – LIVE7 – Confocor 3). The fluorescence of the different reporter proteins was detected using appropriate filters for DAPI, GFP, Texas Red (for mCherry), and Cy7 (for the Fluorescent i-particles NIRFiP-180).

Tissue collection and histological analysis

For histological analysis, treated and control B6D2 testes were fixed by immersion in 4 % paraformaldehyde (PFA) for 14 h. They were then embedded in paraffin before cutting into 5 µm sections using a microtome (Leica biosystems, Wetzlar, Germany). After deparaffination, the sections were stained with hematoxylin and eosin. Stained sections were digitized at 20x magnification using an Axioscan Z1-slide scanner (Zeiss, Jna, Germany). Spermatogenesis was assessed by measuring the area of seminiferous tubules and the cross sections of round tubules (μm2) (n > 35 tubules per testis section; n=5 testis sections per condition). Statistical significance of differences was determined using a Student’s t-test.

Ex vivo Fluorescence Analysis

To analyze the expression of the reporter proteins GFP and mCherry in seminiferous tubules, whole testes were examined under an inverted microscope (CKX53, Olympus, Shinjuku, Tokyo, Japan). Exogenous fluorescence was detected using filters for GFP and Texas Red.

Added in the revised version: Armc2 Genotyping

Added in the revised version: DNA was obtained by the digestion of an ear punch using the Direct PCR lysis buffer (Euromedex, Souffelweyersheim, FR) and 1 mg/ml proteinase K (Thermo Fisher Scientific, Waltham, Massachusetts, USA) overnight at 55°C, followed by enzyme inactivation 1 hour at 85°C. PCR was performed as follow: 2 µl of undiluted DNA was mixed with 0.4 µM primers against Armc2 gene (Genscript, Piscataway, NJ, USA, sequences are listed in table 1), Coral Load PCR Buffer 1X (Qiagen, Hilden, DE), 0.8 mM dNTP mix (Thermo Fisher Scientific) and 0.625 units of Taq DNA Polymerase (Qiagen) in a final volume of 25 µL. The PCR program used was 94°C 15 min (94°C 30 s, 59°C 30 s, 72°C 1 min) x 32 followed by a finale elongation (72°C 10 min). PCR products were then visualized by electrophoresis (BioRad, Hercules, California, USA) on a 1% agarose gel with Gelgreen 1,8X (Interchim, Montluçon, FR).

Added in the revised version: Purification of Seminiferous Tubule Cells

Added in the revised version: Depending on the experimental condition, testes were collected at 1, 3-, 7-, 14-, or 21-days post-injection of mRNA-GFP (300 ng/µL) and transferred into a sterile 35 mm Petri dish containing 2 mL of 1× Phosphate Buffered Saline (PBS, pH 7.4). The tunica albuginea was carefully removed under sterile conditions. Testicular tissue was then mechanically dissociated by gently aspirating and expelling the tissue 10–12 times using an 18-gauge needle to isolate individual seminiferous tubules. To enrich for GFP-positive tubules, tubules lacking GFP fluorescence were discarded under a fluorescence microscope. This procedure resulted in a sample containing approximately 25% GFP-positive tubules.

Added in the revised version: The isolated tubules were washed three times with PBS to remove residual debris. Enzymatic dissociation was performed in two steps: first, the tubules were incubated in collagenase (1 mg/mL in RPMI 1640 medium (Sigma-Aldrich, Saint-Louis, Missouri, USA)) for 30 minutes at 37 °C with gentle agitation. This was followed by a second digestion step using trypsin (2 mg/mL in RPMI 1640 medium (Sigma-Aldrich, Saint-Louis, Missouri, USA)) for 15 minutes at 37°C with gentle agitation to ensure complete dissociation of seminiferous epithelial cells.

Added in the revised version: To stop the enzymatic digestion, cells were immediately resuspended in RPMI 1640 medium (Sigma-Aldrich, Saint-Louis, Missouri, USA) supplemented with 5% KnockOut™ Serum Replacement (KSR; Thermo Fisher Scientific). The resulting single-cell suspension was then used for downstream applications.

Added in the revised version: Live mRNA-GFP Fluorescence Imaging of Dissociated Seminiferous Tubule Cells

Added in the revised version: The isolated seminiferous tubule cell suspensions obtained from testes 24 hours post-injection with mRNA-GFP (300 ng/µL) were cultured in RPMI 1640 medium (Sigma-Aldrich, Saint-Louis, Missouri, USA) in chambered coverglass systems (Nunc™ Lab-Tek™ II, Cat. No. 155360PK, Thermo Fisher Scientific, Courtaboeuf, France) at 37 °C in a humidified atmosphere containing 5% CO₂.

Added in the revised version: After 12 hours of culture, the cells were counterstained with 1.8 µM Hoechst 33342 (Thermo Fisher Scientific), a membrane-permeable nuclear dye functionally similar to DAPI. Live-cell fluorescence imaging was performed using an Axio Observer 7 inverted fluorescence microscope (Zeiss, Jena, Germany) equipped with a 63× oil immersion objective. GFP and Hoechst fluorescence signals were detected using the appropriate filter sets. Image acquisition and analysis were conducted using ZEN lite software (Zeiss, Jena, Germany).

Added in the revised version: Cell Culture

Added in the revised version: HEK cells were cultured in 100 mm tissue culture dishes (Falcon, Corning, New York, USA) in DMEM/F12 medium (Sigma-Aldrich, Saint Louis, Missouri, USA) supplemented with 10% heat-inactivated fetal bovine serum (FBS; Life Technologies), 1% penicillin/streptomycin (Sigma-Aldrich), and 0.01 mg/mL hygromycin B (Thermo Fisher Scientific). Cells were maintained at 37°C in a humidified atmosphere containing 5% CO₂ and were passaged twice weekly at a 1:10 dilution.

Added in the revised version: Transfection was performed using 2 µg/mL GFP mRNA (TriLink BioTechnologies, San Diego, USA) and the jetMESSENGER® mRNA transfection reagent (Sartorius, Göttingen, Germany), following the manufacturer’s instructions. mRNA was extracted from the cells at 12-, 24-, 48-, and 60-hours post-transfection.

Added in the revised version: mRNA Extraction and qPCR Analysis

Added in the revised version: Total RNA from isolated seminiferous tubule cell suspensions and HEK cells was extracted using the NucleoSpin® RNA Plus kit (Macherey-Nagel, Düren, Germany) according to the manufacturer’s protocol. RNA concentration was measured using the Qubit™ RNA Assay Kit (Life Technologies, Carlsbad, California, USA), and samples were stored at –80°C until further use.

Added in the revised version: Up to 800 ng of total RNA per sample was reverse-transcribed into cDNA using the iScript™ cDNA Synthesis Kit (Bio-Rad, Hercules, California, USA), following the manufacturer’s protocol. Gene expression was analyzed by quantitative PCR (qPCR) using 1 µL of undiluted cDNA in a final reaction volume of 10 µL, with gene-specific primers (see Table 1), and the SsoAdvanced™ Universal SYBR® Green Supermix (Bio-Rad, Hercules, California, USA).

Added in the revised version: The qPCR program for GAPDH expression in HEK cells included an initial denaturation at 95°C for 3 minutes, followed by 39 cycles of 95°C for 10 seconds, 60°C for 30 seconds, and 72°C for 30 seconds. Melt curve analysis was performed by increasing the temperature from 58°C to 95°C in 0.5°C increments, holding each step for 2–5 seconds.

Added in the revised version: For GFP expression analysis, the same program was used, except that the annealing temperature was set to 62°C instead of 60°C.

Added in the revised version: Relative gene expression was calculated using the 2^–ΔCT method.

Harris-Shorr sperm Analysis

Earlier version: Sperm were collected from the caudae epididymides of mice (Control, injected with EEV-GFP, GFP-mRNA, or Armc2-mRNA). They were allowed to swim for 10 min at 37 °C in 1 mL M2 media. Sperm were centrifuged at 500 g, washed once with PBS, and fixed 4 % paraformaldehyde in PBS for 1 min at RT. After washing with 1 mL acetate ammonia (100 mM), the sperm suspensions were spotted onto 0.1 % poly L-lysine precoated slides (Thermo Scientific, Courtaboeuf, France) and left to dry. Harris–Schorr staining was then performed according to the WHO protocol [24], and at least 150 sperm were analyzed per animal.

Revised version: Sperm were collected from the caudae epididymides of mice (Control, injected with EEV-GFP, GFP-mRNA, or Armc2-mRNA). They were allowed to swim for 10 min at 37 °C in 1 mL M2 media. Sperm were centrifuged at 500 g, washed once with PBS, and fixed 4 % paraformaldehyde in PBS for 1 min at RT. After washing with 1 mL acetate ammonia (100 mM), the sperm suspensions were spotted onto 0.1 % poly L-lysine precoated slides (Thermo Scientific, Courtaboeuf, France) and left to dry. Harris–Schorr staining was then performed according to the WHO protocol [31], and at least 150 sperm were analyzed per animal.

Whole Testis Optical clearing and 3D image reconstructions

Optical clearing (adapted from uDISCO and Fast 3D clear protocols)

Adult mice were euthanized by cervical dislocation and then transcardiac perfused with 1X PBS (Sigma Aldrich, Saint-Quentin-Fallavier, France). The testes were extracted and fixed for two days at 4 °C in 4 % PFA (Electron Microscopy Sciences, Hatfield, PA, USA). Samples were then washed with PBS for at least two days. Mouse testes were subsequently dehydrated in graded series of tert-butanol solutions (Sigma Aldrich, Saint-Quentin-Fallavier, France) at 35 °C as follows: 30 % overnight (O/N), 50 % for 24 h, 70 % for 10 h, 80 % O/N, 90 % for 10 h, 96 % O/N, and 100 % for 10 h. The testes were cleared in clearing solution (96 % Histodenz, 2 % Sorbitol, 20 % Urea) for 2 days. Then, nuclei were stained with 3.6 μM DAPI (20 % DMSO; 2 % Triton, 1 % BSA) for 2 days. Finally, the testes were then conserved in the clearing solution at 4°C until observation by microscopy. All these steps were carried out under agitation and protected from light.

3D tissue Imaging

The optically-cleared mouse testes were imaged on a lightsheet fluorescence microscope (Blaze, Miltenyi Biotec, Germany), using a 4x NA 0.35 MI PLAN objective protected by a dipping cap for organic solvents, with an overall working distance of 15 mm. Acquisitions on the horizontal plane were obtained with a fixed lightsheet thickness of 3.9 µm at both 488 nm and 561 nm with no overlap between horizontal planes. Voxel resolution x = 1.21; y = 1.21; z = 2 μm. 3D reconstructions were created using Imaris software (Oxford Instruments plc, Abingdon, UK).

Cellular image analysis

Earlier version: The optically-cleared mouse testes were imaged using a ‘ConfoBright’ system which is a unique adaptive optics confocal microscope (Nikon A1R MP, Nikon Europe B.V., The Netherlands) equipped with a deformable mirror module (AOS-micro, AlpAO, Montbonnot, France) to correct geometrical aberrations. Indeed, the 3D confocal imaging of cleared sample requires long distance objectives for deep tissue imaging, resulting in optical aberrations due to inhomogeneous refractive index. The ‘ConfoBright’ microscope corrects these geometric optical aberrations both in excitation and detection light path and restores locally the high diffraction-limited spatial resolution and the best possible photon collection efficiency. The driving metrics based on molecular brightness was used to optimize the adaptive optics in the isoplanatic patch of ca. 100 mm at each imaged depth. Images were acquired using an apo LWD 40x/1.15 water immersion objective (WD 600 µm) and a Plan apo 10x/0.45 objective. FIJI software (open-source software[25]) was used to process and analyze images and Imaris software for the 3D reconstructions.

Revised version: The optically-cleared mouse testes were imaged using a ‘ConfoBright’ system which is a unique adaptive optics confocal microscope (Nikon A1R MP, Nikon Europe B.V., The Netherlands) equipped with a deformable mirror module (AOS-micro, AlpAO, Montbonnot, France) to correct geometrical aberrations. Indeed, the 3D confocal imaging of cleared sample requires long distance objectives for deep tissue imaging, resulting in optical aberrations due to inhomogeneous refractive index. The ‘ConfoBright’ microscope corrects these geometric optical aberrations both in excitation and detection light path and restores locally the high diffraction-limited spatial resolution and the best possible photon collection efficiency. The driving metrics based on molecular brightness was used to optimize the adaptive optics in the isoplanatic patch of ca. 100 mm at each imaged depth. Images were acquired using an apo LWD 40x/1.15 water immersion objective (WD 600 µm) and a Plan apo 10x/0.45 objective. FIJI software (open-source software[32]) was used to process and analyze images and Imaris software for the 3D reconstructions.

Bioluminescence imaging

In vivo Bioluminescence imaging was performed at several time points after in vivo Luciferase -mRNA or EEV-GFP-luciferase injection and electroporation (n=5 mice per condition). Ten minutes before imaging, mice received an intraperitoneal injection of 150 µg g−1 of D-luciferin (Promega, Charbonnières France), and were then anesthetized (isoflurane 4 % for induction and 1.5 % for maintenance) before placing in the optical imaging system (IVIS Lumina III, PerkinElmer, Courtaboeuf, France). In vivo bioluminescence signals were measured in selected regions of interest (injected testes) and were expressed as mean photons per second (mean ± SEM). Background bioluminescence was measured on images from control mice. When the in vivo bioluminescence signal was no longer detectable, testes were collected and immersed in a luciferin solution for a few minutes before performing ex vivo imaging to confirm the absence of signal.

Sperm motility

The cauda epididymis was dilacerated in 1 mL of M2 medium (Sigma Aldrich, Saint-Quentin-Fallavier, France) and spermatozoa were allowed to swim out for 10 min at 37 °C. After incubation, 30 µl of the sperm suspension was immediately placed onto analysis chamber (2X-CEL Slides, 80 µm depth, Leja Products B.V., The Netherlands) kept to 37°C for microscopic quantitative study of sperm movement. Motility of the spermatozoa was evaluated at 37°C with an Olympus microscope and Computer Aided Sperm Analysis (CASA) (CEROS II apparatus; Hamilton Thorne, Beverley, MA, USA). The settings employed for analysis were: acquisition rate: 60 Hz; number of frames: 30; minimum contrast: 30; minimum cell size: 4; low-size gate: 0.13; high-size gate: 2.43; low-intensity gate: 0.10; high-intensity gate: 1.52; minimum elongation gate: 5; maximum elongation gate: 100; magnification factor: 0.81.

The motility parameters measured were: straight line velocity (VSL); curvilinear velocity (VCL); averaged path velocity (VAP); amplitude of lateral head displacement (ALH); beat cross frequency (BCF); linearity (LIN); straightness (STR). Hyperactivated sperm were characterized by VCL > 250 µm s−1, VSL > 30 µm s−1, ALH > 10 µm and LIN < 60, Intermediate by VCL > 120 µm s−1 and ALH > 10 µm, progressive sperm by VAP > 50 µm s−1 and STR > 70 % and slow sperm by VAP < 50 µm s−1and VSL < 25 µm s−1.

Western blot

Western blotting was performed on HEK-293T cells (ATCC, Manassas, VA, USA) transfected with EEV-Armc2 or Armc2-mRNA. Cells were transfected using JetPrime (101000027; Polyplus Illkirch, France) for DNA and JetMessenger (101000056; Polyplus Illkirch, France) for mRNA vectors, both according to the supplier’s recommendations. After 48 h, the cells were washed with PBS and scraped off before centrifuging at 4 °C, 1500 RPM for 5 min. The cell pellet was then resuspended in lysis buffer (87787; Thermo Scientific, Courtaboeuf, France) supplemented with an EDTA-free cocktail of protease inhibitors (11836170001; Roche, Bale, Swiss). The suspension was stirred at 4 °C for 2 h, and then centrifuged at 16,000 g at 4 °C for 10 min. The protein content of the supernatant was estimated with QuantiPro™ BCA Assay kit (Sigma Aldrich, Saint-Quentin-Fallavier, France) before adding 5X Laemmli + 5 % β-mercaptoethanol and heating at 95 °C for 10 min. For the Western blot, 30 μg of proteins were deposited on a ready-made Bis-Tris gel 12 % (Thermo Fisher, Courtaboeuf, France). After the transfer, the PVDF membrane was blocked with 5 % milk in Tris-Buffered Saline solution containing 0.1 % Tween 20 (TTBS) before immunodetection. The anti HA antibody (11867423001; Sigma Aldrich, Saint-Quentin-Fallavier, France) was diluted in TTBS at 1/5000. After incubation with secondary antibodies (AP136P; Sigma Aldrich, Saint-Quentin-Fallavier, France) diluted at 1:10,000 in TTBS, binding was revealed with an enhanced chemiluminescence detection kit (1705062; Clarity Max Western ECL Substrate; BIORAD, Marnes-la-Coquette, France). Membranes were imaged on a ChemiDoc™ system (BIORAD, Marnes-la-Coquette, France).

Immunofluorescence

Immunofluorescence analysis of dissociated testicular cells was performed as follows. After collection, the tunica albuginea was removed from the testes. Then the tissue was mechanically separated with 18G needles. Once washed with PBS, the testicular cells were placed in a dissociation medium containing 1 mg collagenase type V /mL RPMI for 20 min at 37 °C. After filtration (100 μm filter) and centrifugation (5 min at 200 g) the pellets were resuspended in PBS before centrifugation once again. The pellet was then fixed in 1 mL PFA 4 % for 5 min. Then 10 mL of ammonium acetate (0.1 M) was added. Finally, 2 mL of medium was spread on a slide. Testicular cells were permeabilized with 0.1 % PBS-Triton X-100 for 20 min at room temperature. Slides were then blocked in 10 % BSA with PBS-Tween 0.1 % for 30 min before incubating overnight at 4 °C with primary antibodies anti-rabbit ARMC2 (1/50) (HPA053696, Sigma Aldrich, Saint-Quentin-Fallavier, France) and anti-guinea pig tubulin (1/100) (AA345, Swiss Institute of Bioinformatics, Lausanne, Swiss) diluted in PBS-Tween 0.1 % - 5 % BSA. Slides were washed with PBS before incubating for 1 h with secondary antibodies: anti-guinea pig (1/500) (A-11073, Thermo Fischer, Courtaboeuf, France) and anti-rabbit (1:1000) (A-11036, Thermo Fischer, Courtaboeuf, France). Samples were counterstained with DAPI and mounted with DAKO mounting media (NC2313308; Life Technology, Courtaboeuf, France). Fluorescence images were acquired under a confocal microscope (Zeiss, Jena, Germany) fitted with a 63× oil immersion objective. Images were analyzed with ZEN lite software (Zeiss, Jena, Germany).

Intra-cytoplasmic sperm injection (ICSI)

Collection of gametes for ICSI

Armc2−/− sperm or Armc2−/− rescued sperm were harvested by dilaceration of the cauda epididymis. They were allowed to swim for 10 min at 37°C in 1 ml of CZB.HEPES medium containing (in mM) (HEPES 20, NaCl 81.6, KCl 4.8, MgSO4 1.2, CaCl2 1.7, KH2PO4 1.2, EDTA.Na2 0.1, Na-lactate 31, NaHCO3 5, Na-pyruvate 0.3, polyvinyl alcohol 0.1 mg mL−1, phenol red 10 mg mL−1 (0.5 % (w/v) in DPBS), pH 7.4) KCl 125, NaCl 2.6, Na2HPO4 7.8, KH2PO4 1.4 and EDTA 3 (pH 7.0)).The sperm head was separated from the tail by applying multiple piezo pulses (PiezoXpert®, Eppendorf, Montesson, France)

Eggs preparation

CD1 female mice, 7 weeks old, were superovulated by IP injection of 7.5 IU pregnant mare’s serum gonadotrophin (PMSG; Centravet, Dinan, France) followed by 7.5 IU HCG (Centravet, Dinan, France) 48 h later. Eggs were collected from oviducts about 14 h after HCG injection. Cumulus cells were removed with 0.1 % hyaluronidase in M2 medium for 5–10 min. Eggs were rinsed thoroughly and kept in KSOM medium containing (in g L−1: NaCl 5.55, KCl 0.19, KH2PO4 0.05, MgSO4 0.05, NaHCO3 2.1, CaCl2 ⋅ 2H2O 0.250, EDTA 0.004, L-Glutamine 0.146, Na-lactate 1.870, Na-pyruvate 0.020, Glucose 0.04, Penicillin 0.05, Streptomycin 0.07, BSA 1.000, NEAA 0.5 % and EAA 1 %.), at 15°C for at least 15 min until required for ICSI.

ICSI procedures

Earlier version: ICSI was performed according to the method described by Yoshida and Perry (2007)[26].For microinjection, Armc2-/- sperm or Armc2-/- -rescued motile sperm heads were separated from the tail by applying multiple piezo pulses (PiezoXpert®, Eppendorf, Montesson, France). Sperm heads were introduced into the ooplasm using micromanipulators (Micromanipulator InjectMan®, Eppendorf, Montesson, France) mounted on an inverted Nikon TMD microscope (Nikon, Minato-ku, Tokyo, Japan). Eggs that survived the ICSI procedure were incubated in KSOM medium at 37 °C under in an atmosphere of 5 % CO2. Pronucleus formation was checked at 6 h after ICSI, and outcomes were scored up to the blastocyst stage.

Revised version: ICSI was performed according to the method described by Yoshida and Perry (2007)[33].For microinjection, Armc2−/− sperm or Armc2−/− -rescued motile sperm heads were separated from the tail by applying multiple piezo pulses (PiezoXpert®, Eppendorf, Montesson, France). Sperm heads were introduced into the ooplasm using micromanipulators (Micromanipulator InjectMan®, Eppendorf, Montesson, France) mounted on an inverted Nikon TMD microscope (Nikon, Minato-ku, Tokyo, Japan). Eggs that survived the ICSI procedure were incubated in KSOM medium at 37 °C under in an atmosphere of 5 % CO2. Pronucleus formation was checked at 6 h after ICSI, and outcomes were scored up to the blastocyst stage.

In vitro fertilization (IVF)

Earlier version: Armc2-/- sperm or Armc2-/- rescued sperm were harvested by dilaceration of the cauda epididymis. They were allowed to swim in IVF well in capacitated media (M16 + 4 % BSA ) at 37 °C for 10 minutes. Eggs were collected from mature CD1 females (6 weeks old) synchronized with 7.5 units of pregnant mare serum gonadotrophin (PMSG) and 7.5 units of human chorionic gonadotrophin (hCG) prior to collection. Cumulus were incubated for 10 minutes in 500 µL M16 (MR-016; Sigma Aldrich, Saint-Quentin-Fallavier, France) / 1mg L-1 collagenase (C8176, Sigma Aldrich). Cumulus-free and zona-free eggs were collected, and rinsed twice with 500 µl M16 medium. Eggs were incubated with either Armc2-/- sperm or Armc2-/- rescued sperm) in capacitated medium (37 °C, 5 % CO2) for 4 hours. After incubation, unbound sperm were washed away and eggs were incubated with KSOM at 37 °C, 5 % CO2. Twenty-four hours after fertilization, we scored the different stages (unfertilized oocytes, aborted embryos, and 2-cell embryos as an indication of successful fertilization).

Revised version: WT B6D2 sperm, Armc2−/− sperm or Armc2−/− rescued sperm were harvested by dilaceration of the cauda epididymis. They were allowed to swim in IVF well in capacitated media (M16 + 2 % BSA) at 37 °C for 10 minutes. Eggs were collected from mature CD1 females (6 weeks old) synchronized with 7.5 units of pregnant mare serum gonadotrophin (PMSG) and 7.5 units of human chorionic gonadotrophin (hCG) prior to collection. Cumulus were incubated for 10 minutes in 500 µL M16 (MR-016; Sigma Aldrich, Saint-Quentin-Fallavier, France) / 1mg L−1 collagenase (C8176, Sigma Aldrich). Cumulus-free and zona-free eggs were collected, and rinsed twice with 500 µl M16 medium. Eggs were incubated with either WT B6D2 sperm, Armc2−/− sperm or Armc2−/− rescued sperm in capacitated medium (37 °C, 5 % CO2) for 4 hours. After incubation, unbound sperm were washed away and eggs were incubated with KSOM at 37 °C, 5 % CO2. Twenty-four hours after fertilization, we scored the different stages (unfertilized oocytes, aborted embryos, and 2-cell embryos as an indication of successful fertilization).

Statistical analyses

Statistical analyses were performed using SigmaPlot (version 10; Systat Software, Inc., San Jose, CA, USA). To account for sample variability between animals, a paired t-test, Mann-whitney rank sum test, on-way Anova and Wilcoxon test were used. Data are displayed as mean ± SEM. P values of * ≤ 0.05, ** ≤ 0.01, or *** ≤ 0.001 were considered to represent statistically significant differences.

Results

In vivo microinjection and electroporation of mouse testes

Two routes have been described for microinjection of DNA into the testes: direct injection through the tunica albuginea, or injection into the lumen of the seminiferous tubules via the rete testis. We chose the rete testis route and evaluated the efficacy of the microinjection protocol. In particular, we wished to better characterize the diffusion of the injected solution in the volume of the testis, as we were unable to find any information on this parameter in the literature. The efficacy of microinjection via rete testis was assessed using fluorescent i-particles NIRFiP-180, and by measuring their diffusion in testis cross sections examined by microscopy 3 days post-injection (Fig 1). To avoid lesions due to overfilling, the injection was controlled by measuring the expansion of the staining of the peripheral seminiferous tubules during the injection. Injections were stopped when the testes were filled to 2/3 of their capacity (Fig 1A-B). In testis cross sections, the fluorescent i-particles NIRFiP-180 were heterogeneously distributed, and mainly observed in seminiferous tubules located in the peripheral region of the testes, with fewer particles present in the center of the testes (Fig 1C-D). Moreover, no fluorescent i-particles NIRFiP-180 were visible in the peritubular space. These results indicated that microinjection through the rete testis did not produce a homogenous distribution of the particles throughout the seminiferous tubules. Nevertheless, the seminiferous tubules remained intact, as no signal was observed in the peritubular space (Fig 1C-D).

Removed in the revised version: Distribution of I-particles NIRFIP-180 in testis injected via the rete testis route.

Removed in the revised version: (A) A solution containing 0.05 % Fast Green and 1 % fluorescent i-particles NIRFiP-180 was prepared, 10 µL was injected into the seminiferous tubules of adult males, through the rete testes and its efferent channels. Injection was performed at constant pressure under a binocular microscope. The progression of filling of the seminiferous tubules was monitored thanks to the Fast Green. (B) The testes were only filled to 2/3 capacity in order to prevent damage to the tissue. (C) Representative distribution of fluorescent i-particles NIRFiP-180 in a whole cross-section of an injected testis. Nuclei were counterstained with DAPI (blue emission) to reveal tubules. (D) Enlargement of a seminiferous tubule showing particles localized inside the lumens of the tubules. Scales bars: 1 mm and 500 µm.

Next, we assessed the overall safety of the rete testis microinjection and electroporation of mRNA and EEV into testes. The safety of the protocol was evaluated by comparing macroscopic and microscopic anatomies of control (injected with control solution, PBS, 0.05 % FG), and treated testes (injected either with EEV-GFP (PBS, 0.05 % FG) or GFP-mRNA (PBS, 0.05 % FG)). Three days post-injection, the testes were first observed under a binocular microscope to identify possible macroscopic degeneration of the seminiferous tubules (Fig 2 A1 and 2 B1). Degenerations appears as pearly white lesions at the surface of the testis as illustrated in Supp Fig 2 following over electroporation. With the protocol we have developed, no such lesions were observed. Next, the testes were measured and weighed. No statistical differences in length and weight were observed between control and treated testes (Fig 2 A2, A3, B2, B3). Then, microscopic differences were sought by histological analysis of 5 µm sections (Fig 2C). No difference was observed between the control condition and EEV-GFP or GFP-mRNA on the full cross sections (Fig 2 C1, C2, C3). Next, we observed all the different testicular cells, including all germ cell types and Sertoli cells (Fig 2 D1, D2, D3) for each condition. The layered structure of germ cells was identical in all conditions. Analysis of the histological sections revealed no differences in the tubules area of the testes injected either with EEV-GFP or GFP-mRNA (Fig 2 E). At last, Harris-Shorr staining of the epididymal sperm cells demonstrated that there were no increases in morphological defects when mRNA and EEV were used in comparison with the control (Fig 2 F4). Taken together, these results suggest that in vivo microinjection and electroporation of EEV or mRNA did not perturb spermatogenesis.

In vivo injection and electroporation do not alter the morphological structure of the testes, seminiferous tubules, or sperm cells.

Earlier version: (A, B) Testicular morphology was not affected by in vivo injection and electroporation of EEV- GFP (A) or GFP-mRNA (B). Controls correspond to contralateral testes injected/electroporated with control solution (PBS, 0.05 % FG). (A1, B1) comparison of the testicular morphology of adult testes injected with nucleic acid vectors or control solutions. (A2, B2) Comparison of testicular weight and (A3, B3) testicular length on day 7 after injection/electroporation. Data are represented as a box plot median (n=4 for each condition). A Wilcoxon matched pairs test was used to assess the significance of any differences in testis weights and lengths, and p values of ≤ 0.05 were considered statistically significant. (C) Intact testicular structure after in vivo injection and electroporation with EEV-GFP and GFP-mRNA. Comparison of testicular cross section structures. Testes paraffin sections were stained with eosin/hematoxylin and observed by light microscopy (20x magnification). (C1) Control, (C2) EEV-GFP injected and (C3) GFP-mRNA injected. Scales bars: 1000 µm.

Revised version: (A, B) Testicular morphology was not affected by in vivo injection and electroporation of EEV-GFP (A) or GFP-mRNA (B). Controls correspond to contralateral testes injected/electroporated with control solution (PBS, 0.05 % FG). (A1, B1) comparison of the testicular morphology of adult testes injected with nucleic acid vectors or control solutions. (A2, B2) Comparison of testicular weight and (A3, B3) testicular length on day 7 after injection/electroporation. Data are represented as a box plot median (n=4 for each condition). A Wilcoxon matched pairs test was used to assess the significance of any differences in testis weights and lengths, and p values of ≤ 0.05 were considered statistically significant. (C) Intact testicular structure after in vivo injection and electroporation with EEV-GFP and GFP-mRNA. Comparison of testicular cross section structures. Testes paraffin sections were stained with eosin/hematoxylin and observed by light microscopy (20x magnification). (C1) Control, (C2) EEV-GFP injected and (C3) GFP-mRNA injected. Scales bars: 1000 µm. (D) Seminiferous tubule structures are not affected by in vivo injection and electroporation with EEV-GFP and GFP-mRNA. Enlargement of cross sections showing the fine structure of a seminiferous tubule for control (D1), EEV-GFP (D2) and GFP-mRNA (D3). In each tubule the different layers of spermatogenic cells are indicated, Sertoli cells (S), Spermatogonia (Sg), Spermatocytes (Scytes), rounds Spermatids (Stids) mature spermatids cells (m-Sptids), Leydig cells (L). Scales bars: 20 µm. (E) The area of seminiferous tubules is not affected by in vivo injection and electroporation with EEV-GFP and GFP-mRNA. Comparison of the seminiferous tubule diameter after injection of nucleic acid vectors or control solutions. Data are represented as a box plot median. The areas of seminiferous tubules (μm2) were measured for round cross sections of n > 35 tubules per testis section (n= 5 testis sections per condition). Statistical significance was verified using a student’s t-test. (F) Injection/electroporation do not impact epidydimal sperm cells. Representative sperm observed by light microscopy on day 7 after injection/electroporation with Control solution (F1), EEV-GFP (F2), or GFP-mRNA (F3). Scale bars: 10 μm. (F4) Percentage of normal epidydimal sperm cells in each condition. The number of males were n=5 for EEV-GFP; n=6 for GFP-mRNA and n=9 for WT. More than 150 sperm by males were analyzed. Statistical significance was verified using a one-way ANOVA test.

Removed in the revised version: (D) Seminiferous tubule structures are not affected by in vivo injection and electroporation with EEV-GFP and GFP-mRNA. Enlargement of cross sections showing the fine structure of a seminiferous tubule for control (D1), EEV-GFP (D2) and GFP-mRNA (D3). In each tubule the different layers of spermatogenic cells are indicated, Sertoli cells (S), Spermatogonia (Sg), Spermatocytes (Scytes), rounds Spermatids (Stids) and sperm cells (Spz), Leydig cells (L). Scales bars: 20 µm.

Removed in the revised version: (E) The area of seminiferous tubules is not affected by in vivo injection and electroporation with EEV-GFP and GFP-mRNA. Comparison of the seminiferous tubule diameter after injection of nucleic acid vectors or control solutions. Data are represented as a box plot median. The areas of seminiferous tubules (μm2) were measured for round cross sections of n > 35 tubules per testis section (n= 5 testis sections per condition). Statistical significance was verified using a Student’s t-test.

Removed in the revised version: (F) Injection/electroporation do not impact epidydimal sperm cells. Representative sperm observed by light microscopy on day 7 after injection/electroporation with Control solution (F1), EEV-GFP (F2), or GFP-mRNA (F3). Scale bars: 10 μm. (F4) Percentage of normal epidydimal sperm cells in each condition. The number of males were n=5 for EEV-GFP; n=6 for GFP-mRNA and n=9 for WT. More than 150 sperm by males were analyzed. Statistical significance was verified using a one-way ANOVA test.

Analysis of EEV-GFP and GFP-mRNA testicular expression by whole testis imaging

Earlier version: After validating the injection method, we compared the kinetics of GFP expression and the maintenance of the fluorescent signals for mRNA and EEV. To do so, we injected and electroporated one testis of adult B6D2 mice with EEV-GFP (n=129) or with GFP-mRNA (n=65). At 0-, 1-, 7-, 15-, 21-, 28-, 35-, 42- and 49-days post-injection, the whole testes were observed under an inverted microscope. The exogenous fluorescence was directly visible at the surface of the testes when illuminated with light at the appropriate wavelength (Fig 3 and Fig 4). No testicular lesions were observed on the testes at any post injection time (Fig 3 A1-H1 and Fig 4 A1-F1). In addition, both GFP-mRNA and EEV-GFP induced GFP expression in the testes (Fig 3 A2-H2 and Fig 4 A2-F2). It is worth noting that both vectors induced GFP expression at one day post-injection. However, the duration of fluorescent signals were different. For EEV, GFP fluorescence was still observable on day 42 for 100 % of samples, and 56 % of samples were positive on day 49, indicating that expression lasted around 1.5 months (Fig 4G). In contrast, for mRNA, 100 % of testes were labeled on day 21, but none showed any fluorescence on day 28 (Fig 4G). Thus, EEV transfection allowed a considerably longer duration of expression than mRNA. (Fig 3 and 4). It is important to underline that the signal measured is the fluorescence emitted by the GFP. This signal is dependent on both the half-lives of the plasmid/mRNA and the GFP. Therefore, the kinetic of the signal persistence (which is called here expression) is a combination of the persistence of the vector and the synthesized protein. In addition to differences in duration of expression, the GFP expression patterns were clearly different: mRNA produced a large, diffuse pattern, highlighting the shape of the seminiferous tubules; EEV-GFP produced a punctiform pattern (Fig 3B and 4B.) These results suggest that GFP-mRNA and EEV-GFP targeted different seminiferous cell types, such as Sertoli cells and all germline cells, or that there were differences in terms of transfection efficiency.

Revised version: After validating the injection method, we compared the kinetics of GFP expression and the maintenance of the fluorescent signals for mRNA and EEV. To do so, we injected and electroporated one testis of adult B6D2 mice with EEV-GFP (n=129) or with GFP-mRNA (n=65). At 0-, 1-, 7-, 15-, 21-, 28-, 35-, 42- and 49-days post-injection, the whole testes were observed under an inverted microscope. The exogenous fluorescence was directly visible at the surface of the testes when illuminated with light at the appropriate wavelength (Fig 3 and Fig 4). No testicular lesions were observed on the testes at any post injection time (Fig 3 A1-H1 and Fig 4 A1-F1). In addition, both GFP-mRNA and EEV-GFP induced GFP expression in the testes (Fig 3 A2-H2 and Fig 4 A2-F2). It is worth noting that both vectors induced GFP expression at one day post-injection. However, the duration of fluorescent signals were different. For EEV, GFP fluorescence was still observable on day 42 for 100 % of samples, and 56 % of samples were positive on day 49, indicating that expression lasted around 1.5 months (Fig 4G). In contrast, for mRNA, 100 % of testes were labeled on day 21, but none showed any fluorescence on day 28 (Fig 4G). Thus, EEV transfection allowed a considerably longer duration of expression than mRNA. (Fig 3 and 4). It is important to underline that the signal measured is the fluorescence emitted by the GFP. This signal is dependent on both the half-lives of the plasmid/mRNA and the GFP. Therefore, the kinetic of the signal persistence (which is called here expression) is a combination of the persistence of the vector and the synthesized protein. In addition to differences in duration of expression, the GFP expression patterns were clearly different: mRNA produced a large, diffuse pattern, highlighting the shape of the seminiferous tubules; EEV-GFP produced a punctiform pattern (Fig 3B and 4B.)

Kinetics of EEV-GFP expression following in vivo injection/electroporation: whole testicular expression

(A1-H1) Whole-mount testes on days 0, 1, 7, 14, 21, 28, 35 and 42 after in vivo injection/electroporation with EEV-GFP. (A2-H2) Under fluorescence observation, GFP expression was detectable in transfected testes from 12-week-old B6D2 mice. (C3-H3) Insets show the absence of autofluorescence in non-transfected control testes, observed under 4X magnification. The GFP expression presented a punctiform pattern in seminiferous tubules and was detected from 1 day to 42 days. Scales bars: 1 mm and 100 μm.

Kinetics of GFP-mRNA expression following in vivo injection/electroporation: whole testicular expression

Earlier version: (A1-F1) Whole-mount testes on days 0, 1, 7, 15, 21 and 28 after in vivo injection/ electroporation with GFP-mRNA. (A2-F2) Under fluorescence observation, GFP expression was detectable in transfected testes from 12-week-old B6D2 mice. (A3-F3) Insets show the absence of autofluorescence in non-transfected control testes, observed under 4X magnification. The GFP expression presented a continuous pattern in seminiferous tubules and was detected from day 1 to day 15. Scale bars: 1 mm and 100 μm.

Revised version: (A1-F1) Whole-mount testes on days 0, 1, 7, 15, 21 and 28 after in vivo injection/ electroporation with GFP-mRNA. (A2-F2) Under fluorescence observation, GFP expression was detectable in transfected testes from 12-week-old B6D2 mice. (A3-F3) Insets show the absence of autofluorescence in non-transfected control testes, observed under 4X magnification. The GFP expression presented a continuous pattern in seminiferous tubules and was detected from day 1 to day 15. Scale bars: 1 mm and 100 μm. (G) Comparison of the percentage of injected mice exhibiting reporter gene expression. Mice injected with GFP-mRNA exhibited GFP expression from day 1 to day 21. Mice injected with EEV-GFP exhibited GFP expression from day 1 to day 49. (for EEV-GFP n=11 on day 1; n=13 on day 2; n=10 on day 3; n=14 on day 7; n= 5 on day 10; n= 12 on day 15; n=11 on day 21; n= 12 on day 28; n=15 on day 35; n=17 on day 42 and n=9 on day 49); ( for GFP-mRNA n=3 on day 1; n=4 on day 3; n=15 on day 7; n= 21 on day 15; n=15 on day 21 and n= 5 on day 28).

Added in the revised version: These results suggest that GFP-mRNA and EEV-GFP targeted different seminiferous cell types, such as Sertoli cells and all germline cells, or that there were differences in terms of transfection efficiency. Moreover, the stability of mRNA-GFP was assessed by RT-QPCR in HEK cells and seminiferous tubule cells (Fig. 5). mRNA-GFP was detected for up to 60 hours in HEK cells and for up to two weeks in seminiferous tubule cells (Fig. 5). Together, these results suggest that the long-lasting fluorescence observed in our experiments reflects a combination of transcript stability, local translation within germ cells and the slow protein turnover that is typical of the spermatogenic lineage.

Added in the revised version: Stability of mRNA-GFP in HEK cells and seminiferous tubule cells.

Added in the revised version: (A–C) Analysis of mRNA-GFP stability in HEK cells. (A) Quantification of relative GFP mRNA levels (2^-ΔCT) at 24-, 48-, and 60-hours post-transfection. (B) RT-qPCR amplification curves of GFP and (C) GAPDH transcripts at the indicated time points, showing a progressive decrease in GFP signal over time while GAPDH expression remained constant. (D–F) Analysis of mRNA-GFP stability in mouse seminiferous tubule cells after in vivo injection and electroporation. (D) Quantification of relative GFP mRNA levels (2^-ΔCT) at 0-1-, 7-, and 14-days post-transfection. (E) RT-qPCR amplification curves of GFP and (F) ACTB transcripts at corresponding time points. mRNA-GFP remained detectable for up to two weeks in seminiferous tubules, indicating enhanced transcript stability in the testicular environment.

Removed in the revised version: (G) Comparison of the percentage of injected mice exhibiting reporter gene expression. Mice injected with GFP-mRNA exhibited GFP expression from day 1 to day 21. Mice injected with EEV-GFP exhibited GFP expression from day 1 to day 49. (for EEV-GFP n=11 on day 1; n=13 on day 2; n=10 on day 3; n=14 on day 7; n= 5 on day 10; n= 12 on day 15; n=11 on day 21; n= 12 on day 28; n=15 on day 35; n=17 on day 42 and n=9 on day 49) ; ( for GFP-mRNA n=3 on day 1; n=4 on day 3; n=15 on day 7; n= 21 on day 15; n=15 on day 21 and n= 5 on day 28).

Kinetics of EEV and mRNA expression assessed by in vivo imaging

Earlier version: To further assess and compare the kinetics of the expression of the two vectors, we expressed exogenous luciferase in the testis using EEV or mRNA and observed the level of luciferase by in vivo bioluminescence imaging. For EEV, we took advantage of the fact that the EEV plasmid contains the DNA sequence of the luciferase protein (CAGs-GFP-T2A-Luciferase) in addition to the DNA sequence of the GFP fluorescent protein (Supp Fig 1). For mRNA, we injected commercial naked luciferase-mRNA into the rete testis according to the same protocol as used for GFP-mRNA. For this set of experiments, we injected the EEV-GFP-Luc and luciferase-mRNA into the testes of 6 adult mice on day 0. We injected a similar number of copies of mRNA and EEV. The testes were imaged in vivo to detect bioluminescence expression, on day 1 and weekly until disappearance of the signal, no more than 120 days post-injection (Fig 5). For EEV-GFP-Luciferase, the bioluminescence induced by transfection was detected from day 1. After a rapid decrease in signal intensity over the first 3 weeks, a weak but constant signal remained detectable for 3 months, then faded away (Fig 5 A1-A2). For Luciferase-mRNA, expression was also detected from day 1. The bioluminescence decreased gradually over 3 weeks, becoming undetectable thereafter (Fig 5B1-B2). These results are consistent with our previous results (Fig 3) and confirm that EEV allows a longer expression of exogenous protein within the testis. Fig 5C compare the kinetics of expression observed with EEV and mRNA. Overall, our results indicated a stronger expression for mRNA than for EEV, but with expression decreasing rapidly over-time, and almost no remaining signal after 3 weeks. In contrast, residual expression was detectable for several months with EEV.

Revised version: To further assess and compare the kinetics of the expression of the two vectors, we expressed exogenous luciferase in the testis using EEV or mRNA and observed the level of luciferase by in vivo bioluminescence imaging. For EEV, we took advantage of the fact that the EEV plasmid contains the DNA sequence of the luciferase protein (CAGs-GFP-T2A-Luciferase) in addition to the DNA sequence of the GFP fluorescent protein (Supp Fig 1). For mRNA, we injected commercial naked luciferase-mRNA into the rete testis according to the same protocol as used for GFP-mRNA. For this set of experiments, we injected the EEV-GFP-Luc and luciferase-mRNA into the testes of 6 adult mice on day 0. We injected a similar number of copies of mRNA and EEV. The testes were imaged in vivo to detect bioluminescence expression, on day 1 and weekly until disappearance of the signal, no more than 120 days post-injection (Fig 6). For EEV-GFP-Luciferase, the bioluminescence induced by transfection was detected from day 1. After a rapid decrease in signal intensity over the first 3 weeks, a weak but constant signal remained detectable for 3 months, then faded away (Fig 6 A1-A2). For Luciferase-mRNA, expression was also detected from day 1. The bioluminescence decreased gradually over 3 weeks, becoming undetectable thereafter (Fig 6B1-B2). These results are consistent with our previous results (Fig 3) and confirm that EEV allows a longer expression of exogenous protein within the testis. Fig 6C compare the kinetics of expression observed with EEV and mRNA. Overall, our results indicated a stronger expression for mRNA than for EEV, but with expression decreasing rapidly over-time, and almost no remaining signal after 3 weeks. In contrast, residual expression was detectable for several months with EEV.

Kinetics of EEV and mRNA expression by in vivo bioluminescence imaging.

Earlier version: (A) In vivo bioluminescence imaging quantification of luciferase expression over time following injection/electroporation of EEV-GFP-luc. (A1) EEV-GFP-Luc was injected into the testes and electroporated on day 0. Bioluminescence signal was quantified at several time points. Results are expressed as a percentage of the maximal signal (mean ± SEM; n=5 mice up to D2; n=4 from D3 to D28; n=3 from D35 to D98 and n=3 from D105 to D119). (A2) In vivo bioluminescence images of a representative mouse at several time points after administering EEV-GFP-LUC or PBS, and ex vivo bioluminescence images of testes after 119 days.

Revised version: (A) In vivo bioluminescence imaging quantification of luciferase expression over time following injection/electroporation of EEV-GFP-luc. (A1) EEV-GFP-Luc was injected into the testes and electroporated on day 0. Bioluminescence signal was quantified at several time points. Results are expressed as a percentage of the maximal signal (mean ± SEM; n=5 mice up to D2; n=4 from D3 to D28; n=3 from D35 to D98 and n=3 from D105 to D119). (A2) In vivo bioluminescence images of a representative mouse at several time points after administering EEV-GFP-LUC or PBS, and ex vivo bioluminescence images of testes after 119 days. (B) In vivo bioluminescence imaging quantification of luciferase expression over time induced by injection/electroporation of LUC-mRNA. (B1) LUC-mRNA was injected into the testes and electroporated on day 0. Bioluminescence signal was quantified in the whole testis at several time points. Results are expressed as a percentage of the maximal signal (mean ± SEM; n = 5 mice). (B2) In vivo bioluminescence images of a representative mouse at several time points after administering LUC-mRNA or PBS, and ex vivo bioluminescence images of caput, testes, and cauda after 28 days. (C) Decay over time of the number of mice expressing reporter genes. Mice were injected on day 0 with LUC-mRNA or EEV-GFP-LUC and the number of mice showing bioluminescence in the testis was counted at different time points, from day 1 to day 119. For EEV-GFP: n=13 at D1; n=13 at D2; n=4 from D3 to D28; n=3 from D35 to D98 and n=3 from D105 to D119. For LUC-mRNA: n = 5 mice for all-time points.

Removed in the revised version: (B) In vivo bioluminescence imaging quantification of luciferase expression over time induced by injection/electroporation of LUC-mRNA. (B1) LUC-mRNA was injected into the testes and electroporated on day 0. Bioluminescence signal was quantified in the whole testis at several time points. Results are expressed as a percentage of the maximal signal (mean ± SEM; n = 5 mice). (B2) In vivo bioluminescence images of a representative mouse at several time points after administering LUC-mRNA or PBS, and ex vivo bioluminescence images of caput, testes, and cauda after 28 days.

Removed in the revised version: (C) Decay over time of the number of mice expressing reporter genes. Mice were injected on day 0 with LUC-mRNA or EEV-GFP-LUC and the number of mice showing bioluminescence in the testis was counted at different time points, from day 1 to day 119. For EEV-GFP: n=13 at D1; n=13 at D2; n=4 from D3 to D28; n=3 from D35 to D98 and n=3 from D105 to D119. For LUC-mRNA: n = 5 mice for all-time points.

Assessing testicular and cellular GFP-mRNA expression using whole testicle optical clearing, lightsheet microscopy, and 3D reconstructions

Earlier version: To better characterize the spatial distribution of GFP-mRNA expression in the testis, we performed whole testicular optical clearing. On day 0, we injected and electroporated testes with GFP-mRNA (n=6). On day 1, we harvested the testes and cleared them with a modified optical clearing protocol, as described in the MM section. After complete clearing (Fig 6 A), we imaged the whole testes using a lightsheet microscope and performed 3D reconstruction from the images stack obtained (videos 1, 2 and Fig 6 B). From this 3D reconstruction, we determined the volume of the testis stained with GFP by measuring the GFP-positive area in each image and multiplying it by the thickness of the z-step (10 µm). Due to optical issues, only half part of the whole testis was acquired, then the sample was turned by 180 ° to acquire the other half part. A total GFP-stained volume of 0.51 mm3 and 0.23 mm3 was determined from the face A and the face B, respectively. The corresponding total volume of half part of the testis was measured as 60 mm3, therefore 0.81 % and 0.24 % of the testis was transfected for the face A and the face B, respectively (Fig 6 B).

Revised version: To better characterize the spatial distribution of GFP-mRNA expression in the testis, we performed whole testicular optical clearing. On day 0, we injected and electroporated testes with GFP-mRNA (n=6). On day 1, we harvested the testes and cleared them with a modified optical clearing protocol, as described in the MM section. After complete clearing (Fig 7 A), we imaged the whole testes using a lightsheet microscope and performed 3D reconstruction from the images stack obtained (videos 1, 2 and Fig 7 B). From this 3D reconstruction, we determined the volume of the testis stained with GFP by measuring the GFP-positive area in each image and multiplying it by the thickness of the z-step (10 µm). Due to optical issues, only half part of the whole testis was acquired, then the sample was turned by 180 ° to acquire the other half part. A total GFP-stained volume of 0.51 mm3 and 0.23 mm3 was determined from the face A and the face B, respectively. The corresponding total volume of half part of the testis was measured as 60 mm3, therefore 0.81 % and 0.24 % of the testis was transfected for the face A and the face B, respectively (Fig 7 B).

Testicular and cellular GFP-mRNA expression measured on optically cleared testis after 3D image reconstructions from lightsheet microscopy imaging.

Testes were injected/electroporated with GFP-mRNA on day 0. On day 1, whole testes were fixed and subjected to optical clearing. (A) Testes were observed before and after optical clearing on a binocular microscope. The right image shows the transparency of the testis after complete clearing, revealing the blue mesh throughout the organ. (B) The 3D internal structure of a cleared testis was reconstructed from the lightsheet microscopy images. The reconstruction was possible only for a half testis due to optical issues. Two opposing faces of the same testis are presented, allowing the distribution of GFP fluorescence throughout the seminiferous tubules to be measured. Pink fluorescence corresponds to the autofluorescence of interstitial cells located around the seminiferous tubules. Scale bars A: 1 mm and B: 500 µm.

Assessing GFP cellular expression using whole testicle optical clearing and adaptive optics confocal microscopy

Earlier version: Because the GFP fluorescence patterns were different for the two nucleic vectors when observed under the inverted microscope (Fig 3), we wondered whether the same cell types were targeted in both cases. To address this question, the whole optical cleared testes from EEV-GFP and GFP-mRNA-transfected mice were imaged using a confocal microscope. The different cell types were identified based on their positions within the seminiferous tubule, their cellular shape and their nuclear morphology - revealed by nuclear staining. For instance, Sertoli cells have an oval to elongated nucleus and the cytoplasm presents a complex shape (“tombstone” pattern) along the basement membrane [27]. Round spermatids have small, round and compact nuclei with a nucleolus and are localized between the spermatocytes and elongated spermatids [28]. For EEV-GFP, on day 1 post injection and electroporation, a strong punctiform green fluorescence was visible inside the seminiferous tubules (Fig 7 A). Based on the different morphological criteria, this fluorescent signal was detected in spermatocytes, round spermatids and Sertoli cells (Fig 7 BC). On day 7, the GFP signal induced by EEV-GFP was reduced and only isolated signals in a few seminiferous tubules (1 per 11 tubules) were observed (Fig 7 D). These signals were associated only with Sertoli cells only (Fig 7 E).

Revised version: Because the GFP fluorescence patterns were different for the two nucleic vectors when observed under the inverted microscope (Fig 3), we wondered whether the same cell types were targeted in both cases. To address this question, the whole optical cleared testes from EEV-GFP and GFP-mRNA-transfected mice were imaged using a confocal microscope. The different cell types were identified based on their positions within the seminiferous tubule, their cellular shape and their nuclear morphology - revealed by nuclear staining. For instance, Sertoli cells have an oval to elongated nucleus and the cytoplasm presents a complex shape (“tombstone” pattern) along the basement membrane, with long projections that extend toward the lumen. [34]. Round spermatids have small, round and compact nuclei with a nucleolus and are localized between the spermatocytes and elongated spermatids [35]. Depending on the stage of the seminiferous epithelium, round spermatids (steps 1–8) and elongated spermatids (steps 9–16) can be observed in adjacent regions of the same tubule, reflecting the continuous and overlapping progression of spermatid differentiation rather than their constant coexistence within the same stage. For EEV-GFP, on day 1 post injection and electroporation, a strong punctiform green fluorescence was visible inside the seminiferous tubules (Fig 8 A). Based on the different morphological criteria, this fluorescent signal was detected in spermatocytes, round spermatids and Sertoli cells (Fig 8 BC). On day 7, the GFP signal induced by EEV-GFP was reduced and only isolated signals in a few seminiferous tubules (1 per 11 tubules) were observed (Fig 8 D). These signals were associated only with Sertoli cells only (Fig 8 E).

Cellular expression of EEV-GFP following in vivo injection/ electroporation.

Testes were injected/electroporated with EEV-GFP on day 0. On day 1 and on day 7, whole testes were fixed and subjected to optical clearing. Cleared tested were observed by fluorescence microscopy. (A1-A3) On day 1, transfected seminiferous tubules showed dotted green fluorescence at low magnification (10x/0.45). Nuclei were counterstained with DAPI (blue staining) to reveal the structure of the seminiferous tubules. At the cellular level, fluorescence was detectable (B1-B3) in germ cells including Spermatogonia (Sg), Spermatocytes (Scytes) and round Spermatids (RStids), as well as (C1-C3) in Sertoli cells (SC). (D1-D3) On day 7, the GFP signal was lower at low magnification (10x/0.45) and detectable (E1-E3) only in Sertoli cells (40x/1.15 WI) (n=3) (PTc =Peri-tubular myoid cell). E4 is an enlargement of the red square in E3, allowing the cell type to be identified. Scale bars: 100 μm, 15 μm and 3 μm.

Added in the revised version: Germ cells are non-adherent in culture, which allowed us to distinguish them from potential somatic cells that may not have been completely lysed during the dissociation step. Moreover, germ cells have specific sizes and nuclear morphologies. For example, round spermatids are small and possess small, round, compact nuclei with a nucleolus. Nuclear staining with Hoechst allows us to examine the nuclear morphology of cells expressing GFP and, by combining this observation with cell size measurements, to determine the cell type.

Added in the revised version: Thus, 24 hours post-injection and electroporation, we were able to observe that GFP was expressed in spermatocytes, round spermatids, and elongated spermatids

Earlier version: For the mRNA vector, on day 1 and day 7 post-injection and electroporation, an intense fluorescence was observed in all the germ cells and in Sertoli cells in the seminiferous tubules (Fig 8A). At the cellular level, this fluorescent signal was associated with spermatogonia, spermatocytes, round spermatids, elongated spermatids, and mature spermatids cells to similar extents for both post-injection times (Fig 8 B, D). In contrast to what was observed with EEV on day 7, no reduction in the number of fluorescent seminiferous tubules was noted when using GFP-mRNA, with 8 out of 10 tubules stained (Fig 8 CD).

Revised version: For the mRNA vector, on day 1 and day 7 post-injection and electroporation, an intense fluorescence was observed in all the germ cells and in Sertoli cells in the seminiferous tubules (Fig 9A). At the cellular level, this fluorescent signal was associated with spermatogonia, spermatocytes, round spermatids, elongated spermatids, and mature spermatids cells to similar extents for both post-injection times (Fig 9 B, D). In contrast to what was observed with EEV on day 7, no reduction in the number of fluorescent seminiferous tubules was noted when using GFP-mRNA, with 8 out of 10 tubules stained (Fig 9 CD).

Cellular expression of GFP-mRNA following in vivo injection/electroporation.

Earlier version: Testes were injected/electroporated with GFP-mRNA on day 0. On day 1 and day 7, whole testes were fixed and subjected to optical clearing. Cleared testes were observed by fluorescence microscopy. (A1-A3) On day 1, transfected seminiferous tubules showed strong broad-ranging green fluorescence at low magnification (10x0/0.45). Nuclei were counterstained with DAPI (blue staining) to reveal the structure of the seminiferous tubule. At the cellular level, fluorescence was detectable in germ cells (B1-B3) including Spermatogonia (Sg), Spermatocytes (Scytes) and round Spermatids (RStids), mature spermatids cells (m- Sptids) and Sertoli cells (SC). B4 is an enlargement of the red square in B3, allowing the cell types to be identified. (D1-D3) On day 7, the GFP signal remained strong at low magnification (10x/0.45) and was still detectable in (E1-E3) all germ cell types and Sertoli cells (40x/1.15 WI) (n=3). E4 is an enlargement of the red square in E3, showing that testicular sperm were also stained. Scale bars: 100 μm, 15 μm and 3 μm.

Revised version: Testes were injected/electroporated with GFP-mRNA on day 0. On day 1 and day 7, whole testes were fixed and subjected to optical clearing. Cleared testes were observed by fluorescence microscopy. (A1-A3) On day 1, transfected seminiferous tubules showed strong broad-ranging green fluorescence at low magnification (10×0/0.45). Nuclei were counterstained with DAPI (blue staining) to reveal the structure of the seminiferous tubule. At the cellular level, fluorescence was detectable in germ cells (B1-B3) including Spermatogonia (Sg), Spermatocytes (Scytes) and round Spermatids (RStids), mature spermatids cells (m-Sptids) and Sertoli cells (SC). B4 is an enlargement of the red square in B3, allowing the cell types to be identified. (D1-D3) On day 7, the GFP signal remained strong at low magnification (10x/0.45) and was still detectable in (E1-E3) all germ cell types and Sertoli cells (40x/1.15 WI) (n=3). E4 is an enlargement of the red square in E3, showing that testicular sperm were also stained. Scale bars: 100 μm, 15 μm and 3 μm.

Added in the revised version: Assessing GFP-mRNA expression in dissociated germ cells following in vivo testicular electroporation

Added in the revised version: To assess the efficiency of in vivo mRNA delivery into testicular germ cells, GFP-mRNA was injected and electroporated into the testes on day 0. On day 1, testes were collected, enzymatically dissociated, (Fig supp 10) and the resulting seminiferous tubule cell suspensions were cultured for 12 hours. Live cells were then analyzed by fluorescence microscopy (Fig 10). We observed GFP expression in various germ cell types, including pachytene spermatocytes (53,4 %) (Fig 10 A-), round spermatids (25 %) (Fig 10B-E) and in elongated spermatids (11,4%) (Fig 10 C-E). Fluorescence imaging revealed strong cytoplasmic GFP signals in each of these populations, confirming efficient transfection and translation of the delivered mRNA. These results demonstrate that the in vivo injection and electroporation protocol enables effective mRNA transfection across multiple stages of spermatogenesis.

Added in the revised version: GFP-mRNA expression in germ cells following in vivo injection and electroporation.

Added in the revised version: Testes were injected and electroporated with GFP-mRNA on day 0. On day 1, the testes were dissociated, and the resulting seminiferous tubule cell suspension was cultured for 12 hours. Live seminiferous cells were then observed by fluorescence microscopy (LD plan-Neofluar 40x/0.6 Korr M27). (A) Representative images of GFP expression in elongated sperm. Panels show brightfield, DAPI staining (nuclei in blue), GFP fluorescence (green), merged channels, and overlay with brightfield. Magnified views of selected cells are shown below each row, highlighting strong cytoplasmic GFP signals in elongated sperm. (B) Representative images of GFP expression in pachytene spermatocytes. Same panel layout as in (A). Enlarged views show paired pachytene spermatocytes displaying robust cytoplasmic GFP fluorescence. (C) Representative images of GFP expression in round spermatids. Same panel layout as above, showing isolated round spermatids with clear GFP fluorescence in the cytoplasm. Higher magnification confirms successful mRNA transfection and translation. Scale bars: 50 and 5 μm. (D) Percentage of GFP-positive and GFP-negative cells in the analyzed population. (E) Distribution of GFP-positive cells across different germ cell types, showing higher uptake in spermatocytes compared with round and elongated spermatids. (F) Average diameters of spermatocytes and round spermatids, indicating cell size differences in the population analyzed.

Expression of naked mCherry-mRNA in testis following electroporation

Heterologous expression is well known to depend on the protein studied, we therefore tested the same reporter proteins to compare the mRNA and EEV vectors in the experiments presented above. Apart from the bioluminescence experiments with luciferase, we compared only GFP protein expression. To validate and confirm the capacity of naked mRNA to express proteins in the testes after injection and electroporation, we further challenged the method with mCherry, another reporter protein (Supp Fig 1BD). We injected homemade naked mRNA coding for mCherry into the testes. As previously with GFP-mRNA, no testicular lesions were observed (Supp Fig 3 A1, B1, C1, D1, E1, F1).

The assessment of the temporal persistence of testicular mCherry fluorescent protein expression revealed a robust red fluorescence from day 1 post-injection, which remained detectable for at least 15 days (Supp Fig 3 B2, C2, and D2). At the cellular level, the fluorescent signal was detected in germ cells, including spermatogonia, spermatocytes, round spermatids, mature spermatids, and Sertoli cells on days 1 and 7 post-injection (Supp Fig 4).

Finally, we compared the kinetics and levels of expression of the three different mRNA molecules, coding for mCherry, GFP and luciferase. By comparing the number of mice expressing mCherry-mRNA, GFP-mRNA and luciferase-mRNA fluorescence/luminescence over 21 days, we observed first that expression was detectable for all mRNAs on day 1, and second that the duration of expression was slightly different for the individual mRNAs. For instance, on day 15, 100 %, 80 % and 60 % of mice injected with GFP-mRNA, mCherry-mRNA, and mRNA-luciferase, respectively presented fluorescence/bioluminescence and on day 21, 100 % of mice expressed GFP, whereas no signal was observed for mCherry or Luciferase (Supp Fig 5).

Endogenous expression of ARMC2 in germ cells

Before attempting to rescue expression, we felt it was important to better characterize Armc2 expression in healthy germ cells, and in particular to study the timing of expression.

Earlier version: IF was used to determine when ARMC2 protein was detectable during spermatogenesis. For these experiments, dissociated cells from testes were observed to detect the presence of ARMC2 on different spermatogenic cells. ARMC2 was present only in the flagella of the elongated spermatids (Fig 9A and Supp Fig 6A). The specificity of the signal was validated using testicular cells from Armc2 KO mice, where no signal was observed on all spermatogenic cells (Fig. 9B). In transversal sections of WT seminiferous tubules, ARMC2 signal was not present in spermatogonia and spermatocytes (Supp Fig 6B), but detected in spermatid layers.

Revised version: IF was used to determine when ARMC2 protein was detectable during spermatogenesis. For these experiments, dissociated cells from testes were observed to detect the presence of ARMC2 on different spermatogenic cells. ARMC2 was present only in the flagella of the elongated spermatids (Fig 11A and Supp Fig 6A). The specificity of the signal was validated using testicular cells from Armc2 KO mice, where no signal was observed on all spermatogenic cells (Fig. 11B). In transversal sections of WT seminiferous tubules, ARMC2 signal was not present in spermatogonia and spermatocytes (Supp Fig 6B), but detected in spermatid layers.

ARMC2 localization in dissociated testicular cells observed by immunofluorescence.

Cells from WT and Armc2 KO mice were counterstained with Hoechst (A1-B1) and stained with antibodies against tubulin (A2-B2, green signal) and ARMC2 (A3-B3, red signal). (A4-B4) overlay of the different staining. In WT mice, ARMC2 is located in the flagellum of spermatids. In KO mice, no ARMC2 signal (red fluorescence) was observed in any cells.

Earlier version: By analyzing the RNA-seq database produced by Gan’s team [29], we show that the mRNA encoding ARMC2 starts to be expressed at the pachytene spermatocyte stage, then shows a gradual increase at the round spermatid stage, and finally becomes predominantly expressed at the elongated spermatid stage (Supp Fig 6C), a result in agreement with a post-meiotic function of the protein. Finally, we also consistently observed a staining at the base of the manchette of elongating spermatids but we have no explanation for that (Fig 9A3).

Revised version: By analyzing the RNA-seq database produced by Gan’s team [36], we show that the mRNA encoding ARMC2 starts to be expressed at the pachytene spermatocyte stage, then shows a gradual increase at the round spermatid stage, and finally becomes predominantly expressed at the elongated spermatid stage (Supp Fig 6C), a result in agreement with a post-meiotic function of the protein. Finally, we also consistently observed a staining at the base of the manchette of elongating spermatids but we have no explanation for that (Fig 11A3).

In conclusion, the results presented here demonstrate that the ARMC2 protein is expressed and translated at the late stages of spermatogenesis.

Co-injection of Armc2-mRNA and eGFP-mRNA followed by electroporation is safe and induces green fluorescence in the seminiferous tubules

Earlier version: We next tested whether the injection and electroporation of Armc2-mRNA molecules had any deleterious effects on testis anatomy and seminiferous tubule structure. We first verified the quality of Armc2-mRNA synthesis by transfecting HEK cells and performing Western blot (Supp Fig 7). After this validation, we co-injected Armc2-mRNA and eGFP-mRNA into the left testes of mice, using the right testes as untreated controls. eGFP-mRNA was co-injected to verify and monitor transfection efficiency. The testes were observed under a binocular microscope at different times (3-, 6-, 10-, 15-, 21-, 28- and 35-days) after electroporation to identify possible macroscopic degeneration of the seminiferous tubules. No morphological defects were observed in the testes co-injected with Armc2-mRNA and eGFP-mRNA. An example of control and injected testes from day 15 is presented in Fig 10 A1, B1. The testes were also weighed at different times post-injection, and the weight ratios of injected testes to non-injected control testes were determined. For all-time points, this ratio was close to 1 (Fig 10 C), confirming that the method and the mRNAs did not cause any injury at the organ level. Next, under blue light, the efficiency of the transfection was assessed by observing the GFP fluorescence at the surface of the testes. GFP fluorescence was observed on testes injected with Armc2-mRNA and eGFP-mRNA 2 weeks after injection (Fig 10 B2), indicating that the naked mRNA was successfully transfected into testicular cells.

Revised version: We next tested whether the injection and electroporation of Armc2-mRNA molecules had any deleterious effects on testis anatomy and seminiferous tubule structure. We first verified the quality of Armc2-mRNA synthesis by transfecting HEK cells and performing Western blot (Supp Fig 7). After this validation, we co-injected Armc2-mRNA and eGFP-mRNA into the left testes of mice, using the right testes as untreated controls. eGFP-mRNA was co-injected to verify and monitor transfection efficiency. The testes were observed under a binocular microscope at different times (3-, 6-, 10-, 15-, 21-, 28- and 35-days) after electroporation to identify possible macroscopic degeneration of the seminiferous tubules. No morphological defects were observed in the testes co-injected with Armc2-mRNA and eGFP-mRNA. An example of control and injected testes from day 15 is presented in Fig 12 A1, B1. The testes were also weighed at different times post-injection, and the weight ratios of injected testes to non-injected control testes were determined. For all-time points, this ratio was close to 1 (Fig 12 C), confirming that the method and the mRNAs did not cause any injury at the organ level. Next, under blue light, the efficiency of the transfection was assessed by observing the GFP fluorescence at the surface of the testes. GFP fluorescence was observed on testes injected with Armc2-mRNA and eGFP-mRNA 2 weeks after injection (Fig 12 B2), indicating that the naked mRNA was successfully transfected into testicular cells.

In vivo co-injection of Armc2-mRNA and eGFP-mRNA followed by electroporation do not affect testes morphology and weight.

Adult WT mouse testes were injected with a solution containing Armc2-mRNA and eGFP-mRNA. After injection, the testes were electroporated and mice were euthanized two weeks later. (A): Whole testis under white and blue lights on a fluorescence microscope (A1): control testes not injected/electroporated (A2): testes injected with Armc2-mRNA and eGFP-mRNA. eGFP-mRNA was co-injected to follow the transfection efficiency. (B): Ratio of injected/electroporated testis weights to control testis weights at several time points post-injection (3-, 6-, 10-, 15-, 21-, 28- and 35-days post-surgery). n= 1 mouse per time.

Motile sperm cells detected in Armc2 KO mice following Armc2-mRNA injection and electroporation into testes

We then assessed whether the injection of Armc2-mRNA into the testes in Armc2 KO mice restored sperm motility. We examined motility of sperm cells present in the caudal part of the epididymis at different times post-injection (3- to 35-days post-injection). For each condition, between 3 and 16 KO mice were used.

Earlier version: The Armc2 KO model used is known to produce sperm cells with short and irregular flagella that are therefore immotile on day 0. No motile sperm were observed on days 3, 6, 10, 15 or 28 after surgery (Fig 11 A). However, motile sperm cells were found in the epididymis of some Armc2 KO mice at 21- and 35-days post-treatment (Fig 11 A). Indeed, 1 in 3 mice had motile sperm at 21 days post-surgery, rising to 3 in 4 mice at 35 days post-injection. Nevertheless, the number of motile sperm observed remained low: 5.5 % after 21-days and 7.15 % after 35- days post-injection (Fig 11 A1). The sperm motility parameters of Armc2-/--rescued motile sperm were characterized in comparison to those of Armc2+/+ sperm using the computer- assisted semen analysis (CASA) system (Fig 11 A2). These parameters included VAP, VSL, VCL, ALH, BCF, and STR. We have observed significant differences between WT and rescued sperm. In particular, the VSL and LIN parameters are lower for rescued sperm. Next sperm were sorted as progressive, intermediate, hyperactivated, or slow according to motility parameters of motile sperm and recorded from their track (Figure 11A3). The percentage of hyperactivated sperm and the proportion of intermediates in the Armc2-/--rescued motile sperm population were found to be increased in comparison to the control. Videos showing sperm motility in different conditions are available in the online material associated with this article (Videos 3 to 6).

Revised version: The Armc2 KO model used is known to produce sperm cells with short and irregular flagella that are therefore immotile on day 0. No motile sperm were observed on days 3, 6, 10, 15 or 28 after surgery (Fig 13 A). However, motile sperm cells were found in the epididymis of some Armc2 KO mice at 21- and 35-days post-treatment (Fig 13 A). Indeed, 1 in 3 mice had motile sperm at 21 days post-surgery, rising to 3 in 4 mice at 35 days post-injection. Nevertheless, the number of motile sperm observed remained low: 5.5 % after 21-days and 7.15 % after 35-days post-injection (Fig 13 A1). Although the number of motile sperm remained low: 5.5% at 21 days and 7.15% at 35 days post-injection (Fig. 13 A1), this recovery represents a substantial improvement considering the small fraction of germ cells reached by the current electroporation method (Fig 6 B). The sperm motility parameters of Armc2−/−-rescued motile sperm were characterized in comparison to those of Armc2+/+ sperm using the computer-assisted semen analysis (CASA) system (Fig 13 A2). These parameters included VAP, VSL, VCL, ALH, BCF, and STR. We have observed significant differences between WT and rescued sperm. In particular, the VSL and LIN parameters are lower for rescued sperm. Next sperm were sorted as progressive, intermediate, hyperactivated, or slow according to motility parameters of motile sperm and recorded from their track (Figure 13 A3). The percentage of hyperactivated sperm and the proportion of intermediates in the Armc2−/−-rescued motile sperm population were found to be increased in comparison to the control. Videos showing sperm motility in different conditions are available in the online material associated with this article (Videos 3 to 6).

Sperm motility is restored in Armc2 KO mice at 21 and 35 days after injection and electroporation of Armc2-mRNA.

Earlier version: (A) Adult Armc2 KO mouse testes were injected with a solution containing Armc2-mRNA. After injection, the testes were electroporated. At different times (3-, 6-, 10-, 15-, 21-, 28-, and 35-days post-injection), sperm were extracted from the cauda epididymis of the injected testis, and the sample was then examined with a CASA system to identify the percentage of motile spermatozoa (A1). n= 2 for day 15, n= 3 for days 3, 6 and 21, n= 4 for day 10, and n= 5 for days 28 and 35. (A2) Sperm motility parameters of Armc2-/- -rescued sperm in comparison to Armc2+/+ sperm. The motility parameters measured were: averaged path velocity (VAP); straight line velocity (VSL); curvilinear velocity (VCL); amplitude of lateral head displacement (ALH); beat cross frequency (BCF); straightness (STR); linearity (LIN). Black dots: sperm cells from Armc2 null mice, green dots: sperm cells from Armc2 null mice 35 days after injection with Armc2-mRNA. Results are expressed as mean ± SD. (A3) Sperm motility population of Armc2-/- -rescued sperm in comparison to Armc2-/- sperm. Black column: sperm cells from Armc2 null mice, green column: sperm cells from Armc2 null mice 35 days after injection with ARmc2-mRNA. Statistical significance was verified using a Mann-Whitney sum test. Data are displayed as mean ± SEM. P values of * ≤ 0.05, ** ≤ 0.01, or *** ≤ 0.001 were considered to represent statistically significant differences.

Revised version: (A) Adult Armc2 KO mouse testes were injected with a solution containing Armc2-mRNA. After injection, the testes were electroporated. At different times (3-, 6-, 10-, 15-, 21-, 28-, and 35-days post-injection), sperm were extracted from the cauda epididymis of the injected testis, and the sample was then examined with a CASA system to identify the percentage of motile spermatozoa (A1). n= 2 for day 15, n= 3 for days 3, 6 and 21, n= 4 for day 10, and n= 5 for days 28 and 35. (A2) Sperm motility parameters of Armc2−/− -rescued sperm in comparison to Armc2+/+ sperm. The motility parameters measured were: averaged path velocity (VAP); straight line velocity (VSL); curvilinear velocity (VCL); amplitude of lateral head displacement (ALH); beat cross frequency (BCF); straightness (STR); linearity (LIN). Black dots: sperm cells from Armc2 null mice, green dots: sperm cells from Armc2 null mice 35 days after injection with Armc2-mRNA. Results are expressed as mean ± SD. (A3) Sperm motility population of Armc2−/− -rescued sperm in comparison to Armc2−/− sperm. Black column: sperm cells from Armc2 null mice, green column: sperm cells from Armc2 null mice 35 days after injection with ARmc2-mRNA. Statistical significance was verified using a Mann-Whitney sum test. Data are displayed as mean ± SEM. P values of * ≤ 0.05, ** ≤ 0.01, or *** ≤ 0.001 were considered to represent statistically significant differences. (B) Morphology of sperm cells in Armc2 KO mice injected or not with Armc2-mRNA. (B1-B2): microscopic observation of epididymal sperm cells from a mature WT mouse. (B3-B4): epididymal sperm cells from a mature Armc2 KO mouse 35 days after injection/electroporation with Armc2-mRNA. (B5-B6): epididymal sperm cells from a control Armc2 KO male. Normal sperm cells were observed in the injected condition with Armc2-mRNA. (White arrows), Scale bars: 10 µm.

Removed in the revised version: (B) Morphology of sperm cells in Armc2 KO mice injected or not with Armc2-mRNA. (B1-B2): microscopic observation of epididymal sperm cells from a mature WT mouse. (B3-B4): epididymal sperm cells from a mature Armc2 KO mouse 35 days after injection/electroporation with Armc2-mRNA. (B5-B6): epididymal sperm cells from a control Armc2 KO male. Normal sperm cells were observed in the injected condition with Armc2- mRNA. (White arrows), Scale bars: 10 µm.

Removed in the revised version: After verifying motility, we looked at the morphology of the spermatozoa present in the cauda epididymis. Six days after injection of Armc2-mRNA, the cells detected were mostly round cells and abnormal spermatozoa with a short or coiled flagellum measuring between 7 and 20 µm.

Earlier version: The same cell types were observed at 3-, 10-, 15- and 28-days post-surgery. In contrast, the motile sperm detected on days 21 and 35, had a normal morphology with a long flagellum (greater than 100 µm) and a hook-shaped head (Fig 11 B and Supp Fig 8).

Revised version: After verifying motility, we looked at the morphology of the spermatozoa present in the cauda epididymis. Six days after injection of Armc2-mRNA, the cells detected were mostly round cells and abnormal spermatozoa with a short or coiled flagellum measuring between 7 and 20 µm. The same cell types were observed at 3-, 10-, 15- and 28-days post-surgery. In contrast, the motile sperm detected on days 21 and 35, had a normal morphology with a long flagellum (greater than 100 µm) and a hook-shaped head (Fig 13 B and Supp Fig 8).

Motile sperm cells detected in Armc2 KO mice following Armc2-mRNA injection and electroporation into testes are fertile

Earlier version: We subsequently evaluated the efficacy of Armc2-mRNA injection into the testes of Armc2 KO mice in restoring sperm fertility. The fertility outcome was assessed through in vitro fertilization (IVF) and intracytoplasmic sperm injection (ICSI) experiments. The sperm rescued from Armc2-/- mice were capable of successfully fertilizing eggs and producing embryos at the two-cell stage by IVF (Fig 12 A-B-C). Notably, 62.7% of two-cell embryos were obtained with the Armc2-/--rescued sperm by IVF, compared to only 2.67% with the Armc2-/- sperm. Three % of eggs became 2-cell embryos when fertilized with sperm from Armc2-/-, a rate not significantly different to that observed for eggs incubated 24 h without sperm, and likely corresponding to parthenogenesis activation (Fig 12 A1-A2-B).

Revised version: We subsequently evaluated the efficacy of Armc2-mRNA injection into the testes of Armc2 KO mice in restoring sperm fertility. The fertility outcome was assessed through in vitro fertilization (IVF) and intracytoplasmic sperm injection (ICSI) experiments. The sperm rescued from Armc2−/− mice were capable of successfully fertilizing eggs and producing embryos at the two-cell stage by IVF (Fig 14 A-B-C). The 2-cell embryo formation observed after IVF serves as a strong indication of partial functional recovery. Notably, 62.7% of two-cell embryos were obtained with the Armc2−/−-rescued sperm by IVF, compared to only 2.67% with the Armc2−/−sperm. Three % of eggs became 2-cell embryos when fertilized with sperm from Armc2−/−, a rate not significantly different to that observed for eggs incubated 24 h without sperm, and likely corresponding to parthenogenesis activation (Fig 14 A1-A2-A3-A4-B).

Armc2 −/−-rescued sperm can fertilize eggs and produce embryos by IVF.

Earlier version: (A) Illustrations showing eggs/embryos obtained 24 h after egg collection in the following conditions. (A1) W/O sperm (A2) IVF with Armc2-/- sperm from Armc2-/- males (#444) and (A3) IVF with Armc2-/- rescued sperm from Armc2-/- males treated with mRNA-Armc2 (males #399, #389, and #406). Green and red asterisks show 2-cell embryos, red asterisks showing 2-cell embryos obtained by parthenogenesis. White asterisks show unfertilized oocytes or degenerated. (B) Histograms showing the mean percentage ± SD of alive eggs reaching the 2- cell embryo stage at 24 hours after IVF without sperm (n=4), with Armc2-/- sperm (n=4), and with Armc2-/- rescued sperm (n=5). Statistical significance was verified using a one-way ANOVA test. (C) The table presents the data pertaining to the number of two-cell embryos obtained by IVF with Armc2-/- rescued sperm from five different Armc2-/- males treated with mRNA- Armc2 (males #399, #406, #389, #388, and #395). The data is presented in terms of the percentage of two-cell embryos obtained in relation to the total number of eggs collected.

Revised version: (A) Illustrations showing eggs/embryos obtained 24 h after egg collection in the following conditions. (A1) W/O sperm (A2) IVF with Armc2−/− sperm from Armc2−/− males (#444), (A3) IVF with Armc2−/− rescued sperm from Armc2−/− males treated with mRNA-Armc2 (males #399, #389, and #406) and (A4) IVF with WT B6D2 sperm. Green and red asterisks show 2-cell embryos, red asterisks showing 2-cell embryos obtained by parthenogenesis. White asterisks show unfertilized oocytes or degenerated. (B) Histograms showing the mean percentage ± SD of alive eggs reaching the 2-cell embryo stage at 24 hours after IVF without sperm (n=4), with Armc2−/− sperm (n=4), with Armc2−/− rescued sperm (n=5), and with WT B6D2 sperm (n=5). Statistical significance was verified using a one-way ANOVA test. (C) The table presents the data pertaining to the number of two-cell embryos obtained by IVF with Armc2−/− rescued sperm from five different Armc2−/− males treated with mRNA-Armc2 (males #399, #406, #389, #388, and #395). The data is presented in terms of the percentage of two-cell embryos obtained in relation to the total number of eggs collected.

To gain further insight, a comparative analysis of the developmental outcomes of mouse embryos generated by ICSI with spermatozoa from wild-type (WT), Armc2−/− and Armc2−/−treated mice was performed. It should be noted that in case of Armc2−/− treated mice, the ICSI procedure was performed only with the motile Armc2−/−-rescue spermatozoa. The percentage of live injected oocytes that reached the blastocyst embryos was 46% for WT spermatozoa, 25% for Armc2−/−-rescued spermatozoa and 13% for Armc2−/− spermatozoa (Supp Fig.9). The findings indicate that the developmental potential of the embryos was enhanced when Armc2−/−-rescued sperm were utilized as opposed to Armc2−/− sperm.

Earlier version: Overall, these results demonstrate that the Armc2-/--rescue motile spermatozoa can successfully fertilize eggs and produce embryos capable of developing properly into blastocysts.

Revised version: Overall, these results demonstrate that the Armc2−/−-rescue motile spermatozoa can successfully fertilize eggs.

Discussion

The challenge of treating male infertility remains to be addressed. Current assisted reproduction techniques (ARTs) are unable to treat all patients, and alternative strategies need to be developed to meet the legitimate desire to be a father. The aim of this study was to evaluate the potential of naked mRNA as a means to induce an expression of exogenous proteins in male germ cells in a preclinical adult mouse model. Based on previous studies using electroporation, we investigated whether the combination of the injection of naked mRNA and in vivo electroporation could lead to an efficient protein expression in spermatogenic cells. We chose to first study the efficiency of capped and poly(A)-tailed mRNA coding for reporter proteins and compared the results to those obtained with a non-integrative enhanced episomal vector plasmid. No EEV plasmid has ever been tested in the context of infertility treatment before this study.

Earlier version: Using an adult mouse model, we optimized the micro-injection and electroporation method described by Michaelis et al [15]. We show that the microinjection through the rete testis did not provide a homogenous distribution of the particles throughout the seminiferous tubules. Nevertheless, the seminiferous tubules remained intact, with no signal detected in the peritubular space. The peripheral expression observed was due to the close vicinity of cells to the electrodes, and to a peripheral dispersal of the injected solution, as shown by the distribution of the fluorescent i-particles NIRFiP-180. Our results also showed that the combination of injection and electroporation did not perturb spermatogenesis when electric pulses are carefully controlled. Using such a protocol, we were able to induce the expression of 3 reporter proteins, GFP, mCherry and luciferase in the testis by mRNA injection/electroporation.

Revised version: Using an adult mouse model, we optimized the micro-injection and electroporation method described by Michaelis et al [17]. We show that the microinjection through the rete testis did not provide a homogenous distribution of the particles throughout the seminiferous tubules. Nevertheless, the seminiferous tubules remained intact, with no signal detected in the peritubular space. The peripheral expression observed was due to the close vicinity of cells to the electrodes, and to a peripheral dispersal of the injected solution, as shown by the distribution of the fluorescent i-particles NIRFiP-180. Our results also showed that the combination of injection and electroporation did not perturb spermatogenesis when electric pulses are carefully controlled. Using such a protocol, we were able to induce the expression of 3 reporter proteins, GFP, mCherry and luciferase in the testis by mRNA injection/electroporation.

Using whole testicular optical clearing, the reporter proteins were synthesized from the injected mRNA in different types of cells including Sertoli cells, spermatogonia, spermatocytes, round spermatids, and mature spermatids from day 1, and were still detectable after 1 week. These results deserve two comments: first, the expression is very fast and synthesized protein is detectable 24 h injection, second all cell types have the ability to translate mRNAs. Furthermore, the fact that we observed motile sperm at 21 days after injection confirms that spermatids are transfected and that the translation of the protein of interest is possible at this stage. For EEV, we have a similar result at day 1. However, the yields of seminiferous tubules and cellular transfection are lower. In particular, a lower level of transfection of germ cells was observed than with the mRNA. It is worth noting that after one week, the reporter proteins synthesized from injected EEV, were only discernible in the Sertoli cells.

Added in the revised version: Based on whole testes fluorescence and, for the first time, in vivo bioluminescence imaging of testes, we characterized the kinetics of mRNA expression. The signal measured is the fluorescence or the bioluminescence emitted by the GFP or luciferase. This signal is dependent on both the half-lives of the plasmid/mRNA and the proteins. Therefore, the kinetic of the signal persistence is a combination of the persistence of the vector and the synthesized protein. Although naked mRNAs are generally short-lived in somatic cells, several factors likely explain the prolonged detection of reporting proteins. RT-qPCR analysis showed that mRNA-GFP persisted for up to 60 hours in HEK cells and two weeks in seminiferous tubule cells (Fig. 5), indicating enhanced stability in the testicular environment. Codon optimization and chemical modifications (capping, polyadenylation) are known to improve mRNA stability and translation efficiency [37–40].

Earlier version: Based on whole testes fluorescence and, for the first time, in vivo bioluminescence imaging of testes, we characterized the kinetics of mRNA expression. The signal measured is the fluorescence or the bioluminescence emitted by the GFP or luciferase. This signal is dependent on both the half-lives of the plasmid/mRNA and the proteins. Therefore, the kinetic of the signal persistence is a combination of the persistence of the vector and the synthesized protein. For mRNA, it’s difficult to determine the lifespan of our mRNAs because these mRNAs have been modified at different levels including 5’CAP, mRNA body, poly(A)tail, which increase mRNA stability and translation [30]. Nevertheless, their half-lives should not exceed a few days and therefore the fluorescent signal observed, ranging from 15- to 21-days, depending on the molecule being expressed, likely corresponds to the persistence of the protein synthesized during the time window of mRNA expression. The persistence of the reporter proteins is in line with the fact that proteins involved in spermatogenesis exhibit a markedly low turnover rate [31]. This is due to the fact that these proteins are stored within sperm organelles, such as the acrosome, manchette, centrioles or fibrous sheath. These organelles, made during spermiogenesis, remain stable for weeks until the fertilization process occurs because there is no protein synthesis in mature sperm. For example, the Cav3.2 calcium channel is expressed during meiosis at the pachytene stage and contributes to calcium signaling during acrosome reaction [32–35]. When using the EEV vector, expression persisted for longer – up to 119 days – due to the intrinsic property of the EEV plasmid which allows its replication in synchronous manner with the host genome.

Revised version: The persistence of the reporter proteins is in line with the fact that proteins involved in spermatogenesis exhibit a markedly low turnover rate [41]. This is due to the fact that these proteins are stored within sperm organelles, such as the acrosome, manchette, centrioles or fibrous sheath. These organelles, made during spermiogenesis, remain stable for weeks until the fertilization process occurs because there is no protein synthesis in mature sperm. For example, the Cav3.2 calcium channel is expressed during meiosis at the pachytene stage and contributes to calcium signaling during acrosome reaction [42–45].

Added in the revised version: When using the EEV vector, expression persisted for longer – up to 119 days – due to the intrinsic property of the EEV plasmid which allows its replication in synchronous manner with the host genome.

These results suggest that although EEV expression lasted longer, mRNAs, by targeting more efficiently male germ cells and allowing higher transfection yields of seminiferous tubules, could be a more effective and potent tool to express exogenous proteins in germ cells. By expressing a missing protein in the case of male infertility due to monogenic causes, it could be possible to restore failed spermatogenesis and thus to treat infertility.

ARMC2 is expressed in late spermatogenesis stages

Earlier version: We show that ARMC2 was localized in the flagellum of spermatids obtained by enzymatic dissociation. No Armc2 expression was detected in earlier germ cell type lineages like spermatogonia or spermatocytes. These results suggest that the ARMC2 protein is expressed late during spermatogenesis, which explains why motile sperm were found in the cauda epididymis from 3-weeks after injection in our treated mice. Indeed, full spermiogenesis (from round cells to sperm) takes around 15 days [36], and the journey across the epididymis lasts around 8 days, making a total of 3-weeks. Our results also confirm those recently published by Lechtreck et al. [37] from their study of the role of ARMC2 in the Intra-Flagellar Transport (IFT) of radial spokes in Chlamydomonas. They suggested that the transport of the radial spokes along the flagellum involves ARMC2, acting as an IFT adapter [37]. The presence of ARMC2 in the flagella of elongating spermatid supports this hypothesis.

Revised version: We show that ARMC2 was localized in the flagellum of spermatids obtained by enzymatic dissociation. No Armc2 expression was detected in earlier germ cell type lineages like spermatogonia or spermatocytes. These results suggest that the ARMC2 protein is expressed late during spermatogenesis, which explains why motile sperm were found in the cauda epididymis from 3-weeks after injection in our treated mice. Indeed, full spermiogenesis (from round cells to sperm) takes around 15 days [46], and the journey across the epididymis lasts around 8 days, making a total of 3-weeks. Our results also confirm those recently published by Lechtreck et al. [47] from their study of the role of ARMC2 in the Intra-Flagellar Transport (IFT) of radial spokes in Chlamydomonas. They suggested that the transport of the radial spokes along the flagellum involves ARMC2, acting as an IFT adapter [47]. The presence of ARMC2 in the flagella of elongating spermatid supports this hypothesis.

Exogenous Armc2-mRNA expression rescued the motility of oligo-astheno-teratozoospermic sperm

This is the first demonstration that proteins can be expressed in the testis following electroporation with optimized, capped and poly(A)-tailed mRNA.

Our objective was to develop a new targeted therapeutic approach for infertility associated with monogenic defects. The objective of this preclinical study was to ascertain the efficacy of messenger RNA (mRNA) in expressing a missing protein, ARMC2, in a mouse model exhibiting oligo-astheno-teratozospermia due to the absence of Armc2, with the aim of restoring flagellar motility and fertility. Our results strongly suggest that the strategy did not alter spermatogenesis, as injection and electroporation of Armc2-mRNA or EEV-Armc2 had no effect on testicular morphology or weight. More importantly, the technique was effective, with motile sperm cells found in cauda epididymis 3 and 5 weeks after Armc2-mRNA injection into testes from Armc2 KO males. Nevertheless, it should be noted that not all injected mice were efficiently treated. For example, only 87.5% of the treated mice (14 of the 16) exhibited motile sperm after 5 weeks. The absence of motile sperm at 5 weeks may be attributed to the specific types of spermatogenic cells that were transfected during the electroporation phase. It is possible that the transfected cells may differ between individuals, potentially influenced by the injection and the position of the electrodes during electroporation. If the mRNA transfection occurs in a spermatogonia, it may take more than six weeks (including the time required to cross the epididymis) before motile epididymal sperm cells emerge. This potential timeline could explain the absence of motile sperm in some mice at 5 weeks.

Added in the revised version: Due to the quantity of motile sperm obtained, it was not possible to produce offspring through natural mating. Nevertheless, the rescued Armc2−/− sperm displayed normal morphology and motility and were competent for in vitro fertilization (IVF) and intracytoplasmic sperm injection (ICSI). The rescued sperm successfully fertilized oocytes and produced embryos that developed to the two-cell stage following IVF, and to the blastocyst stage after ICSI. these findings support a partial functional recovery of the sperm population. Taken together, these results demonstrate that mRNA electroporation can partially restore sperm function by rescuing motility and enabling early embryonic development. This proof-of-concept supports the feasibility of mRNA-based protein supplementation in spermatogenic cells, paving the way for future therapeutic development.

Earlier version: Due to the quantity of motile sperm obtained, it was not possible to produce offspring through natural mating. However, the Armc2-/--rescued sperm exhibited normal morphology, motility, and in vitro fertility. Indeed, the Armc2-/--rescued motile spermatozoa have successfully fertilized eggs and produced embryos that were capable of developing properly into blastocysts. These results provide compelling evidence that the method can effectively produce fertile sperm. It worth to note that the significant modifications of the CASA parameters observed for rescued sperm did not impact their fertilizing potential. Naked mRNA injection/electroporation is therefore a promising method to treat infertility. In contrast, no motile spermatozoa were found after injection/electroporation of EEV-Armc2, confirming our previous results suggesting that this nucleic tool does not efficiently enter or transfect germ cells.

Revised version: It worth to note that the significant modifications of the CASA parameters observed for rescued sperm did not impact their fertilizing potential. Naked mRNA injection/electroporation is therefore a promising method to treat infertility. In contrast, no motile spermatozoa were found after injection/electroporation of EEV-Armc2, confirming our previous results suggesting that this nucleic tool does not efficiently enter or transfect germ cells.

Earlier version: Despite this success, the transfection rate deserves to be improved to obtain a larger number of sperm cells to produced offspring through natural mating. To increase the testicular transfection rate, encapsulation of mRNA into lipids nanoparticles could be used, as used for Covid vaccination [38]. During the writing of this manuscript, Dong team [39], used a self- amplifying RNA (saRNA) encapsulated in cholesterol-amino-phosphate derived lipid nanoparticle to restore spermatogenesis in infertile mice. They successfully restore the expression of the DNA Meiotic Recombinase 1 (DMC1) [40–42] in Dmc1 KO infertile mice by injecting a self-amplifying RNA-Dmc1 in the testes. saRNA are genetically engineered replicons derived from self-replicating single-stranded RNA viruses [43]. The saRNA contains the alphavirus replicase genes and encodes an RNA-dependent RNA polymerase (RdRP) complex which amplifies synthetic transcripts in situ and the target RNA sequence. The target RNA is expressed at high levels as a separate entity. As a result of their self-replicative activity, saRNAs can be delivered at lower concentrations than conventional mRNA to achieve comparable expression levels [44]. Moreover, saRNA constructs need to be condensed by a cationic carrier into a nanoparticle measuring ∼100 nm to enable their uptake into target cells and protect the saRNA from degradation [45]. Finally, saRNA will amplify the RNA without cellular regulation. For all these reasons, if such a strategy is to be pursued, a potential toxicity effect due to saRNA over expression must be investigated in the testes and progeny. Another difficulty with saRNA relates to the size of the molecular construct. saRNA sequences are large and complex. The length of the sequence RdRP is around 7 kilobases, which often makes the full length of saRNA more than 9 kilobases once the sequence for the protein of interest has been integrated [45]. Dong et al. [39] successfully used their saRNA construct to rescue spermatogenesis failure induced by the absence of the small protein Dmc1 (37 KDa), but it may be more challenging with larger proteins such as the structural proteins involved in OAT, including the 98-kDa ARMC2 [23].

Revised version: Despite this success, the transfection rate deserves to be improved to obtain a larger number of sperm cells to produced offspring through natural mating. To increase the testicular transfection rate, encapsulation of mRNA into lipids nanoparticles could be used, as used for Covid vaccination [48]. During the writing of this manuscript, Dong team [49], used a self-amplifying RNA (saRNA) encapsulated in cholesterol-amino-phosphate derived lipid nanoparticle to restore spermatogenesis in infertile mice. They successfully restore the expression of the DNA Meiotic Recombinase 1 (DMC1) [50–52] in Dmc1 KO infertile mice by injecting a self-amplifying RNA-Dmc1 in the testes. saRNA are genetically engineered replicons derived from self-replicating single-stranded RNA viruses [53]. The saRNA contains the alphavirus replicase genes and encodes an RNA-dependent RNA polymerase (RdRP) complex which amplifies synthetic transcripts in situ and the target RNA sequence. The target RNA is expressed at high levels as a separate entity. As a result of their self-replicative activity, saRNAs can be delivered at lower concentrations than conventional mRNA to achieve comparable expression levels [54]. Moreover, saRNA constructs need to be condensed by a cationic carrier into a nanoparticle measuring ∼100 nm to enable their uptake into target cells and protect the saRNA from degradation [55]. Finally, saRNA will amplify the RNA without cellular regulation. For all these reasons, if such a strategy is to be pursued, a potential toxicity effect due to saRNA over expression must be investigated in the testes and progeny. Another difficulty with saRNA relates to the size of the molecular construct. saRNA sequences are large and complex. The length of the sequence RdRP is around 7 kilobases, which often makes the full length of saRNA more than 9 kilobases once the sequence for the protein of interest has been integrated [55]. Dong et al. [49] successfully used their saRNA construct to rescue spermatogenesis failure induced by the absence of the small protein Dmc1 (37 KDa), but it may be more challenging with larger proteins such as the structural proteins involved in OAT, including the 98-kDa ARMC2 [25].

Our next step will be to assess whether encapsulating Armc2-mRNA in LNP-CAP could allow a larger number of germ cells to be transfected.

Naked mRNA, a new therapeutic strategy to treat severe infertility

Earlier version: Non-obstructive Azoospermia (NOA) and severe oligozoospermia (SO) are the most severe disorders of spermatogenesis and are the most likely to be of genetic origin. NOA is defined by the complete absence of spermatozoa in the ejaculate. Approximately 10–15 % of infertile men have azoospermia, and a further 15 % have SO [46]. For patients with NOA, few clinical solutions are currently available. Generally, testicular sperm extraction is attempted to collect some spermatozoa from the seminiferous tubules, which can then be used for ICSI [47]. When no sperm are retrieved, intra-conjugal conception is impossible. The results of this study strongly suggest that transient mRNA expression of a missing protein in NOA testes by electroporation could be sufficient to produce normal sperm for IVF and obtain embryos.

Revised version: Non-obstructive Azoospermia (NOA) and severe oligozoospermia (SO) are the most severe disorders of spermatogenesis and are the most likely to be of genetic origin. NOA is defined by the complete absence of spermatozoa in the ejaculate. Approximately 10–15 % of infertile men have azoospermia, and a further 15 % have SO [56]. For patients with NOA, few clinical solutions are currently available. Generally, testicular sperm extraction is attempted to collect some spermatozoa from the seminiferous tubules, which can then be used for ICSI [57]. When no sperm are retrieved, intra-conjugal conception is impossible. The results of this study strongly suggest that transient mRNA expression of a missing protein in NOA testes by electroporation could be sufficient to produce normal sperm for IVF and obtain embryos.

Added in the revised version: In conclusion, this work provides the first demonstration that in vivo testicular injection and electroporation of capped and poly(A)-tailed mRNA can efficiently transfect male germ cells and produce proteins with sufficient stability to support spermatogenic progression. Our comprehensive analysis revealed that mRNA electroporation is more efficient than episomal vector delivery, despite shorter persistence of expression. Importantly, Armc2-mRNA electroporation was able to partially rescue sperm motility and fertilizing ability in Armc2−/−mice, representing a crucial first step toward developing mRNA-based strategies to correct monogenic causes of male infertility.

Removed in the revised version: In conclusion, this paper presents the first in vivo testicular injection and electroporation of capped and poly(A)-tailed mRNA, demonstrating that it is an efficient strategy to transfect male germ cells and that the duration of the expression of the resulting proteins is compatible with restoring spermatogenesis. Our comprehensive study revealed mRNA to be more efficient than an episomal vector, despite longer-lived expression in male germ cells with EEV. The difference was linked to EEV achieving a lower rate of seminiferous tubule transfection and a shorter duration of expression in germ cells. Our findings have also demonstrated, for the first time, that sperm motility and fertility can be restored in mice with an oligo-astheno- teratozoospermia phenotype through a technique that combines injection and electroporation of capped and poly(A)-tailed mRNA. The findings presented open new opportunities to develop efficient strategies to treat male infertility with monogenic causes.